Dual transcranial electromagnetic stimulation of the precuneus boosts human long-term memory

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

The authors make a bold claim that a combination of repetitive transcranial magnetic stimulation (intermittent theta burst-iTBS) and transcranial alternating current stimulation (gamma tACS) causes slight improvements in memory in a face/name/profession task.

Strengths:

The idea of stimulating the human brain non-invasively is very attractive because, if it worked, it could lead to a host of interesting applications. The current study aims to evaluate one such exciting application.

Weaknesses:

(1) The title refers to the "precuneus-hippocampus" network. A clear definition of what is meant by this terminology is lacking. More importantly, mechanistic evidence that the precuneus and the hippocampus are involved in the potential effects of stimulation remains unconvincing.

Thank you for the observation. We believe that the evidence collected supports our state relative to the stimulation of the precuneus and the involvement of the hippocampus. In particular, given the existing evidence on TMS methodology and precuneus non-invasive stimulation (see Koch et al., Brain, 2022, Koch et al., Alzheimer's research & therapy, 2025), the computation of the biophysical model with the E-field we produced (see Biophysical modeling and E-field calculation section in the supplementary information), together with the individual identification of the precuneus through the RM (see iTBS+γtACS neuromodulation protocol and MRI data acquisition in the main text), we can reasonably assume that the individually identified PC was stimulated.

As we acknowledged in the Limitations section, we cannot entirely rule out the possibility that our results might also reflect stimulation of more superficial parietal regions adjacent to the precuneus. Nor do we provide direct evidence of microscopic changes in the precuneus following stimulation. However, the results we provide in terms of changes in precuneus oscillatory activity and precuneus-hippocampi connectivity sustain both our thesis of the precuneus stimulation and of hippocampi involvement in the stimulation effects.

Despite this consideration, we agree on the fact that a clear definition of what is meant by the terminology “precuneus-hippocampus network” is lacking. Moreover, since our data and previous evidence sustain the notion of PC stimulation, while this study does not produce direct evidence of the hippocampi stimulation - but only of the effect of the neuromodulation protocol on its connection with the precuneus, we soften the claim in the title. We remove the mention of the precuneus-hippocampus network so that the modified title will be as follows: “Dual transcranial electromagnetic stimulation of the precuneus boosts human long-term memory.”

(2) The question of the extent to which the stimulation approach and the stimulation parameters used in these experiments causes specific and functionally relevant neural effects remains open. Invasive recordings that could address this question remain out of the scope of this non-invasive study. The authors conducted scalp EEG experiments in an attempt to address this question using non-invasive methods. However, the results shown in Fig. 3 are unclear. The results are inconsistently reported in units of microvolts squared in some panels (3A, 3B) and in units of microvolts in other panels (3C). Also, there is insufficient consideration of potential contamination by signal components reflecting eye movements, other muscle artifacts, or another volume-conducted signal reflecting aggregate activity inside the brain.

As you correctly noted, Figure 3 presents results obtained from the TMS–EEG recordings. However, there is no inconsistency regarding the measurement units, as we are referring to two distinct indices: one in the frequency domain—oscillatory power shown in Figures 3A and 3B, expressed in microvolts squared (μV²)—and one in the time domain—the TMS-evoked potential shown in Figure 3C, expressed in microvolts (μV).

Regarding the concern about artifacts, this is an important issue on which our group has a strong expertise, having published well-established, highly cited procedures on how to record and clean TMS-EEG signals (e.g., Casula et al., Clinical Neurophysiology, 2017; Rocchi et al., Brain Stimulation, 2021). In the current study, we adopted a well-established and rigorous approach for both data acquisition and preprocessing. This ensured that the recorded TMS–EEG signals were not contaminated by physiological or electrical artifacts.

As regards the recording procedure, all participants were instructed to fixate on a black cross to minimize eye movements. To avoid auditory-related components caused by the TMS click, we adopted an ad-hoc procedure optimized for TMS-EEG recordings (Rocchi et al., Brain Stimulation, 2021). First, participants were given earphones that continuously played an ad-hoc masking noise composed of white noise mixed with specific time-varying frequencies of the TMS click (Rocchi et al., Brain Stimulation, 2021). The masking noise volume was adjusted to ensure that participants could not detect the TMS click, or as much as tolerated (always below 90 dB). To further reduce the impact of the TMS click on the EEG signal, we placed ear defenders (SNR=30) on top of the earphones. Please see TMS–EEG data acquisition section in the main text.

As regards the offline cleaning process, we applied Independent Component Analysis (INFOMAX-ICA) to the EEG data to identify and remove components associated with muscle activity, eye movements, blinking, and residual TMS-related artifacts, in line with the most recent guidelines on TMS–EEG preprocessing (Hernandez-Pavon et al., Brain Stimulation, 2023). Specifically, for TMS-related muscle artefacts, we strictly followed the criteria based on their scalp topography, spectral content, timing, and amplitude, which we published in a paper focused on this topic (Casula et al., Clinical Neurophysiology, 2017). We add this detail in the TMS–EEG preprocessing and analysis section in the supplementary information (lines 119-120).

(3) Figure 3 indicates "Precuneus oscillatory activity ...", but evidence that the activity presented reflects precuneus activity is lacking. The maps shown at the bottom of Figure 3C suggest that the EEG signals recorded with scalp EEG reflect activity generated across a wide spatial range, with a peak encompassing at least tens of centimeters. Thus, evidence that effects specifically reflect precuneus activity, as the paper's title and text throughout the manuscript suggest, is lacking.

We believe there may have been a misunderstanding. As indicated in the figure caption, panels A and B represent oscillatory activity, whereas panel C displays the TMS-evoked potentials (TEPs). Therefore, the topographical maps mentioned (i.e., those in panel C) did not refer to oscillatory activity, but to differences in TEP amplitude. Specifically, the topographies shown in Figure 3C illustrate statistically significant differences in TEP amplitudes between post-stimulation time points (T1—immediately after stimulation, and T2—20 minutes after stimulation) and the pre-stimulation baseline (T0).

In this figure, we focused our analysis on a cluster of electrodes overlying the individually identified precuneus, capturing EEG responses to single TMS pulses delivered to that target. This approach, widely used in previous literature (e.g., Koch et al., NeuroImage, 2018; Casula et al., Annals of Neurology, 2022; Koch et al., Brain, 2022; Maiella et al., Clinical Neurophysiology, 2024; Koch et al., Alzheimer’s Research & Therapy, 2025), supports the interpretation that the observed responses reflect precuneus-related activity. Furthermore, the wide spatial range change you mention proved to be statistically different only when conducting the TMS-EEG over the precuneus (i.e., administering the TMS single pulse over the precuneus) and not when performing it over the left parietal cortex. We modified the discussion section in the main text to make it more clear (lines 196-199).

“Moreover, we observed specific cortical changes in the posteromedial parietal areas, as evidenced by the whole-brain analysis conducted on TMS-EEG data when performed over the precuneus and the absence of effect when TMS-EEG was performed on the lateral posterior parietal cortex used as a control condition.”

That said, we do not state that the effects observed specifically reflect the precuneus activity; indeed, we think the effect of the stimulation is broader, as discussed in the Discussion section. We rather sustain, in line with the literature (Koch et al., Neuroimage 2018; Koch et al., Brain, 2022; Koch et al., Alzheimer's research & therapy, 2025), the idea that the effects observed are a consequence of the precuneus stimulation by the dual stimulation.

(4) The paper as currently presented (e.g., Figure 3) also lacks rigorous evidence of relevant oscillatory activity. Prior to filtering EEG signals in a particular frequency band, clear evidence of oscillations in the frequency band of interest should be shown (e.g., demonstration of a clear peak that emerges naturally in the frequency range of interest when spectral analysis is applied to "raw" signals). The authors claim that gamma oscillations change because of the stimulation, but a clear peak in the gamma range prior to stimulation is not apparent in the data as currently presented. Thus, the extent to which spectral measurements during stimulation reflect physiological gamma oscillations remains unclear.

If we understand correctly, your concern relates to the lack of a clear gamma peak before neuromodulation, which may suggest uncertainty about the observed changes in gamma oscillatory activity. Is that correct?

First, it is important to underline that the natural frequency typically observed in the precuneus falls within the beta range, not the gamma range (see Rosanova et al., Journal of Neuroscience, 2009; Casula et al., Annals of Neurology, 2022). This explains why a prominent gamma peak is not expected at baseline (T0).

Differently, our neuromodulatory protocol was specifically aimed at boosting gamma oscillatory activity given its well-established role in learning and memory processes (Griffiths & Jensen, Trends in Neurosciences, 2023). Thus, to assess the effect of the neuromodulatory protocol, we compared the oscillatory activity before (T0) and after stimulation (T1 and T2), which showed a clear increase in the gamma band. This effect is visible in the raw oscillatory power plot and is most clearly represented in Figure 3B, where the gamma band emerged as the only frequency range showing significant changes across time points.

(5) Concerns remain regarding the rigor of statistical analyses in the revised manuscript (see also point 8 below). Figure 3B shows an undefined statistical test with p<0.05. The statistical test that was used is not explained. Also, a description of how corrections for multiple comparisons were made is missing. Figures 3A and 3C are not accompanied by statistics, making the results difficult to interpret. For Figure 4C, a claim was made based on a significant p-value for one statistical test and a non-significant p-value in another test. This is a common statistical mistake (see Figure 1 and accompanying discussion in Makin and Orban de Xivry (2019) Science Forum: Ten common statistical mistakes to watch out for when writing or reviewing a manuscript. eLife 8:e48175).

All statistical tests are described in the Statistical Analysis section of the main text. Specifically, to assess cortical oscillation changes in Experiment 3, we conducted repeated-measures ANOVAs with stimulation condition (iTBS+γtACS vs. iTBS+sham-tACS) and time (ΔT1 = T1–T0; ΔT2 = T2–T0) as within-subject factors, for each frequency band. To further explore the effects of stimulation at each time point, we performed paired t-tests with Bonferroni correction for multiple comparisons. A one-tailed hypothesis was adopted, based on our a priori prediction of gamma-band increase derived from previous work (Maiella et al., 2022).

Please note that Figures 3A and 3C are purely descriptive and are therefore not accompanied by statistical tests. Figure 3A shows the full spectral profile across frequencies and conditions, while statistical significance for these data is reported in Figure 3B. Similarly, the upper part of Figure 3C displays the TMS-evoked potential (TEP) in the precuneus, while the statistical comparison of TEP amplitudes across time points is shown in the lower part of Figure 3C.

Regarding Figure 4C and the article you cited, are you referring to the error described as “Interpreting comparisons between two effects without directly comparing them”? If we understand correctly, this refers to the mistake of inferring an effect by observing that a significant result occurs in one condition or group, while the corresponding result in another condition or group is not significant, without directly testing the difference between them.

In the case of Experiment 4, which investigates fMRI effects and is illustrated in Figure 4, we employed a general linear model that explicitly modeled both conditions and time points, allowing for a direct statistical comparison. Therefore, the connectivity effect reported does not fall into the category of the error you mentioned.

Importantly, Figure 4C does not depict the effect of the neuromodulatory protocol itself. Rather, its purpose is to show that, within the real stimulation condition, there is a correlation between the observed effect and the integrity of the bilateral Middle Longitudinal Fasciculus. No conclusions or assumptions were made based on the absence of a significant correlation in the sham condition. However, since it was an exploratory analysis, we decided to soften our claims relative to the neural mechanism in the discussion section of the main text (lines 241-246).

(6) In the second question posed in the original review, I highlighted that it was unclear how such stimulation would produce memory enhancement. The authors replied that, in the absence of mechanisms, there are many other studies that suffer from the same problem. This raises the question of placebo effects. The paper does not sufficiently address or discuss the possibility that any potential stimulation effects may reflect placebo effects.

We agree with the reviewer on the potential role of a placebo effect in our study. For this reason, our experimental study had several stimulation conditions, including a placebo condition, which corresponded to the sham iTBS-sham tACS condition, which did not produce any effect.

(7) The third major concern in the original review was the lack of evidence for a mechanism that is specific to the precuneus. Evidence for specific involvement of the precuneus remains lacking in the revised manuscript. The authors state: "the non-invasive stimulation protocol was applied to an individually identified precuneus for each participant". However, the meaning of this statement is unclear. Specifically, it is unclear how the authors know that they are specifically targeting the precuneus. Without directly recording from the precuneus and directly demonstrating effects, which is outside of the scope of the study, specific involvement of the precuneus seems speculative. Also, it does not seem as though a figure was included in the paper to show how the stimulation protocol specifically targets the precuneus. In their response to the original reviews, the authors state that posterior medial parietal areas are the only regions that show significant differences following the stimulation, but they did not cite a specific figure, or statistics reported in the text, that show this. In any event, posterior medial parietal areas encompass a wide area of the brain, so this would still not provide evidence for an effect specifically involving the precuneus.

We respectfully disagree with the claim that targeting the precuneus in our study is speculative. The statement that “without directly recording from the precuneus and directly demonstrating effects, which is outside the scope of the study, specific involvement of the precuneus seems speculative” would, by that logic, implicitly call into question a large body of cognitive neuroscience research employing non-invasive techniques such as EEG and fMRI.

Our methodological approach—combining MRI-guided stimulation, biophysical modeling, and TMS–EEG—is well established and widely used for targeting and studying the role of specific cortical regions, including the precuneus (e.g., Wang et al., Science, 2014; Koch et al., NeuroImage, 2018; Casula et al., Annals of Neurology, 2022, 2023; Koch et al., Brain, 2022; Maiella et al., Clinical Neurophysiology, 2024; Koch et al., Alzheimer’s Research & Therapy, 2025).

In line with previously published protocols (Santarnecchi et al., Human Brain Mapping, 2018; Özdemir et al., PNAS, 2020; Mantovani et al., Journal of Psychiatric Research, 2021), we identified individual targets (i.e., the precuneus) for each participant based on structural and resting-state functional MRI data (see MRI Data Acquisition and Preprocessing section in the main text). This target was then accurately localized using MRI-guided stereotaxic neuronavigation, ensuring reproducible and anatomically precise stimulation across subjects.

Finally, concerning the last comment about the lack of figures/statistics showing how the stimulation protocol targets the precuneus and the specificity of the effect observed, we would like to let the focus go over:

Figure 3 in the main text, where we show the results of the TME-EEG over the posterior medial parietal areas;

Figure S1 in the supplementary information, which shows with the e-fied simulation how the stimulation protocol targets the brain;

the Precuneus iTBS+γtACS increases gamma oscillatory activity section in the main text results, where we report the results of the statistical analysis of the TMS-EEG conducted over the precuneus and the left posterior parietal cortex, used as a control condition to test for the specificity of the neuromodulation protocol.

(8) Regarding chance levels, it is unfortunate that the authors cannot quantify what chance levels are in the immediate and delayed recall conditions. This makes interpretation of the results challenging. In the immediate and delayed conditions, the authors state that the chance level is 33%. It would be useful to mark this in the figures. If I understand correctly, chance is 33% in Fig. 2A. If this is the case and if I am interpreting the figure correctly:

Gray bars for the sham condition appear to be below chance (~20-25%). Why is this condition associated with an accuracy level that is lower than chance?

Cyan bars and red bars do not appear to be significantly different from chance (i.e., 33%), with red slightly higher than cyan. What statistic was performed to obtain the level of significance indicated in the figure? The highest average value for the red condition appears to be around 35%. More details are needed to fully explain this figure and to support the claims associated with this figure.

The immediate and recall conditions you mention correspond to a free recall task. In this case, the notion of a fixed "chance level" is not straightforward as it would be in recognition or forced-choice paradigms, which is why we did not quantify it at first. I will now try to explain this extensively.

Unlike multiple-choice tasks, where participants select the answer from a limited set of alternatives and the probability of a correct response by chance can be precisely quantified (e.g., 33% in a 3-alternative forced choice), free recall involves the spontaneous retrieval of items from memory without external cues or predefined options. As such, the response range in free recall is essentially unconstrained, encompassing the entire vocabulary of the participant.

Because of this open-ended nature, the probability of correctly recalling a studied item purely by chance is exceedingly low and could be approximated to zero. Also, in our task, participants had to correctly recollect both name and occupation, doubling the possibility of the answers.

This assumption is further supported by the fact that random guesses in free recall are unlikely to match any of the studied items, given the vast number of possible alternatives. As a result, performance above zero can be reasonably interpreted as reflecting genuine memory retrieval, rather than random guessing.

As regards statistics, repeated-measures ANOVAs with stimulation condition as a within-subject factor (i.e., iTBS+γtACS; iTBS+sham-tACS; sham-iTBS+sham-tACS) for each dependent variable (see statistical analysis section in main text).

(9) In the revised version of the paper, the authors did not address concerns associated with the block design (please see question 4d in the original review).

We are sorry for the misunderstanding. We did not address your concerns related to block design since it does not apply to our study. As reported in the paper you mentioned in the original review, block design involves data collection performed in response to different stimuli of a given class presented in succession. If this is the case, it does not correspond to our experimental design since both TMS-EEG and fMRI were conducted in the resting state (i.e., without the presentation of stimuli) on different days according to the different randomized stimulation conditions.

In sum, this study presents an admirable aspirational goal, the notion that a non-invasive stimulation protocol could modulate activity in specific brain regions to enhance memory. However, the evidence presented at the behavioral level and at the mechanistic level (e.g. the putative involvement of specific brain regions) remains unconvincing.

We hope our response will be carefully considered, fostering a constructive exchange and leading to a reassessment of your evaluation.

Reviewer #2 (Public review):

Summary:

The manuscript by Borghi and colleagues provides evidence that the combination of intermittent theta burst TMS stimulation and gamma transcranial alternating current stimulation (γtACS) targeting the precuneus increases long-term associative memory in healthy subjects compared to iTBS alone and sham conditions. Using a rich dataset of TMS-EEG and resting-state functional connectivity (rs-FC) maps and structural MRI data, the authors also provide evidence that dual stimulation increased gamma oscillations and functional connectivity between the precuneus and hippocampus. Enhanced memory performance was linked to increased gamma oscillatory activity and connectivity through white matter tracts.

Strengths:

The combination of personalized repetitive TMS (iTBS) and gamma tACS is a novel approach to targeting the precuneus, and thereby, connected memory-related regions to enhance long-term associative memory. The authors leverage an existing neural mechanism engaged in memory binding, theta-gamma coupling, by applying TMS at theta burst patterns and tACS at gamma frequencies to enhance gamma oscillations. The authors conducted a thorough study that suggests that simultaneous iTBS and gamma tACS could be a powerful approach for enhancing long-term associative memory. The paper was well-written, clear, and concise.

Comments on Revision:

I thank the authors for their thoughtful responses to my first review and their inclusion of more detailed methodological discussion of their rationale for the stimulation protocol conditions and timing. Regarding the apparent difference in connectivity at baseline between conditions, the explanation that this is due to intrinsic dynamics, state, or noise implies the baseline is reflecting transient changes in dynamics rather than a true or stable baseline. Based on this, it looks like iTBS solely is significantly greater than the baseline before the iTBS and γtACS condition but maybe not that much lower than post-stimulation period for iTBS and γtACS. A longer baseline period should be used to ensure transient states are not driving baseline levels such that these endogenous fluctuations would average out. This also raises questions about whether the effect of iTBS and γtACS or iTBS alone are dependent on the intrinsic state at the time when stimulation begins. Their additional clarification of memory scoring is helpful but also reveals that the effect of dual iTBS+γtACS specifically on the association between faces and names is just significant. This modest increase in associative memory should be taken into consideration when interpreting these findings.

We thank the reviewer for the feedback. We fully agree that considering baseline dynamics is critical when assessing the neurophysiological and connectivity effects of stimulation protocols.

In Experiments 3 and 4, baseline measurements were specifically included in our design to account for the possibility that intrinsic dynamics, state, or noise could influence the observed effects of neuromodulation. Indeed, if we had compared only post-stimulation connectivity between the real and sham conditions, the effects might have appeared larger. The inclusion of baseline measurements allows us to contextualize and better isolate the neuromodulatory impact by controlling such endogenous fluctuations. Importantly, the fMRI connectivity measurements, which comprise the baseline, are derived from 10-minute BOLD signal acquisitions, which help mitigate the influence of transient fluctuations and provide a quite stable estimate of intrinsic connectivity.

Moreover, regarding the possibility that stimulation effects may depend on the intrinsic state at stimulation onset, we hypothesize that gamma-frequency entrainment induced by tACS could reduce the variability of intrinsic dynamics, promoting a more stable neural state that is favorable for the induction of long-term plasticity.

As regards the memory scoring, we would like to clarify that the significant improvement observed in the dual iTBS+γtACS condition does not pertain solely to the face–name association. Rather, it concerns the more demanding task of recalling the association between face, name, and occupation. While we agree that the observed effect could be considered modest, it is worth noting that it follows from only 3 minutes of stimulation.

Reviewer #3 (Public review):

Summary:

Borghi and colleagues present results from 4 experiments aimed at investigating the effects of dual γtACS and iTBS stimulation of the precuneus on behavioral and neural markers of memory formation. In their first experiment (n = 20), they find that a 3-minute offline (i.e., prior to task completion) stimulation that combines both techniques leads to superior memory recall performance in an associative memory task immediately after learning associations between pictures of faces, names, and occupation, as well as after a 15-minute delay, compared to iTBS alone (+ tACS sham) or no stimulation (sham for both iTBS and tACS). Performance in a second task probing short-term memory was unaffected by the stimulation condition. In a second experiment (n = 10), they show that these effects persist over 24 hours and up to a full week after initial stimulation. A third (n = 14) and fourth (n = 16) experiment were conducted to investigate neural effects of the stimulation protocol. The authors report that, once again, only combined iTBS and γtACS increases gamma oscillatory activity and neural excitability (as measured by concurrent TMS-EEG) specific to the stimulated area at the precuneus compared to a control region, as well as precuneus-hippocampus functional connectivity (measured by resting state MRI), which seemed to be associated with structural white matter integrity of the bilateral middle longitudinal fasciculus (measured by DTI).

Strengths:

Combining non-invasive brain stimulation techniques is a novel, potentially very powerful method to maximize the effects of these kinds of interventions that are usually well-tolerated and thus accepted by patients and healthy participants. It is also very impressive that the stimulation-induced improvements in memory performance resulted from a short (3 min) intervention protocol. If the effects reported here turn out to be as clinically meaningful and generalizable across populations as implied, this approach could represent a promising avenue for treatment of impaired memory functions in many conditions.

Methodologically, this study is expertly done! I don't see any serious issues with the technical setup in any of the experiments. It is also very commendable that the authors conceptually replicated the behavioral effects of experiment 1 in experiment 2 and then conducted two additional experiments to probe the neural mechanisms associated with these effects. This certainly increases the value of the study and the confidence in the results considerably.

The authors used a within-subject approach in their experiments, which increases statistical power and allows for stronger inferences about the tested effects. They also used to individualize stimulation locations and intensities, which should further optimize the signal-to-noise ratio.

Weaknesses:

I think one of the major weaknesses of this study is the overall low sample size in all of the experiments (between n = 10 and n = 20). This is, as I mentioned when discussing the strengths of the study, partly mitigated by the within-subject design and individualized stimulation parameters. The authors mention that they performed a power analysis but this analysis seemed to be based on electrophysiological readouts similar to those obtained in experiment 3. It is thus unclear whether the other experiments were sufficiently powered to reliably detect the behavioral effects of interest. In the revised manuscript, the authors provide post-hoc sensitivity analyses that help contextualize the strength of the findings.

While the authors went to great lengths trying to probe the neural changes likely associated with the memory improvement after stimulation, it is impossible from their data to causally relate the findings from experiments 3 and 4 to the behavioral effects in experiments 1 and 2. This is acknowledged by the authors and there are good methodological reasons for why TMS-EEG and fMRI had to be collected in separate experiments, but readers should keep in mind that this limits inferences about how exactly dual iTBS and γtACS of the precuneus modulate learning and memory.

We thank the reviewer for the feedback.

Reviewer #1 (Recommendations for the authors):

I suggest:

(1) Removing all mechanistic claims about the precuneus and hippocampus.

We soften our claims about the precuneus-hippocampus network.

(2) Repeating and focusing on the behavioral experiments with a much larger number of images and stronger statistical power to try to demonstrate a compelling behavioral correlate of the proposed stimulation protocol.

We clarified the misunderstanding relative to the chance level of the behavioral experiments raised by the reviewer.

Reviewer #2 (Recommendations for the authors):

Use longer baseline to establish stable gamma level for comparisons in Figure 3

If we understand correctly, you propose to increase the baseline to establish the gamma oscillatory activity as expressed in Figure 3 (showing the results of experiment 3). Is that right? In the figure, you see a baseline of -100; 0ms, which we use for a merely graphical reason, since no activity is usually observable before the TMS pulse. However, to establish the level of gamma, we used a larger baseline correction ranging from -700 ms to -300 ms (i.e., 400ms). We added this important information in the cortical oscillation section of the supplementary information (lines 134-135).

Reviewer #3 (Recommendations for the authors):

I think that the authors did a great job responding to the concerns raised by the reviewers. All of my own comments have been satisfactorily addressed. I will update my public review to be more concise, so that it only includes the overall assessment of the manuscript, including the strengths and weaknesses, but without the requests for clarification. Strengths and weaknesses remain largely the same, as the authors did not conduct additional experiments.

Thank you.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Non-invasive brain stimulation techniques have the potential to improve memory functions. However, the results so far have been relatively modest and time-consuming. Here, we implemented a novel 3-minute combination of personalized repetitive transcranial magnetic stimulation (intermittent theta burst-iTBS) coupled with simultaneous application of gamma transcranial alternating current stimulation (γtACS) over the precuneus, a brain area connected with the hippocampus, to modulate long term memory in healthy subjects. Only dual electromagnetic stimulation of the precuneus produced a consistent increase in long-term associative memory as compared to iTBS alone and sham conditions in a sample of healthy volunteers. The effects were replicated in another independent sample, in which the increased associative memory was retained for up to one week. Moreover, dual stimulation increased gamma oscillations and precuneus-hippocampus functional connectivity through the white matter tracts linking the precuneus with the temporal lobe. These findings show that dual stimulation may lead neuronal assemblies in a state favorable to enhance long-term plasticity and identify the precuneus as a key brain area involved in memory formation. Personalized dual electromagnetic stimulation of the precuneus-hippocampus network may represent a new powerful approach for enhancing memory functions in several healthy and clinical conditions.

Revised version: Non-invasive brain stimulation techniques have the potential to improve memory functions. However, the results so far have been relatively modest and time-consuming. Here, we implemented a novel 3-minute combination of personalized repetitive transcranial magnetic stimulation (intermittent theta burst-iTBS) coupled with simultaneous application of gamma transcranial alternating current stimulation (γtACS) over the precuneus, a brain area connected with the hippocampus, to modulate long-term memory in healthy subjects. Only dual electromagnetic stimulation of the precuneus produced an increase in long-term associative memory as compared to iTBS alone and sham conditions in a sample of healthy volunteers. The effects were replicated in another independent sample, in which the increased associative memory was retained for up to one week. Moreover, dual stimulation increased gamma oscillations and precuneus-hippocampus functional connectivity through the white matter tracts linking the precuneus with the temporal lobe. These findings show that dual stimulation may lead neuronal assemblies to a state favorable to enhancing long-term plasticity. Personalized dual electromagnetic stimulation of the precuneus may represent a new powerful approach for enhancing memory functions in several healthy and clinical conditions.

Introduction

As the world’s population ages, the prevalence of age-related memory deficits is growing. This phenomenon often signals the onset of more severe cognitive decline, which has a strong impact on society. In light of this escalating challenge, there is an urgent need for innovative strategies aimed at enhancing cognitive functions and potentially mitigating cognitive decline.

Non-invasive brain stimulation (NIBS) techniques have been largely used to enhance human cognition in the last two decades (Antal et al., 2022). Encouraging results have been reported in both healthy and pathological conditions including depression, stroke, and Alzheimer’s disease (Blumberger et al., 2018; Koch et al., 2022; Lefaucheur et al., 2020).

Earlier version: Repetitive transcranial magnetic stimulation (rTMS) and transcranial alternating current stimulation (tACS) are two forms of NIBS widely used to enhance memory performances (Grover et al., 2022; Koch et al., 2018; Wang et al., 2014). rTMS, based on the principle of Faraday, induces depolarization of cortical neuronal assemblies and leads to after-effects that have been linked to changes in synaptic plasticity involving mechanisms of long-term potentiation (LTP) (Huang et al., 2017; Jannati et al., 2023). On the other hand, tACS causes rhythmic fluctuations in neuronal membrane potentials, which can bias spike timing, leading to an entrainment of the neural activity (Wischnewski et al., 2023). In particular, the induction of gamma oscillatory activity has been proposed to play an important role in a type of LTP known as spike timing-dependent plasticity, which depends on a precise temporal delay between the firing of a presynaptic and a postsynaptic neuron (Griffiths and Jensen, 2023). Both LTP and gamma oscillations have a strong link with memory processes such as encoding (Bliss and Collingridge, 1993; Griffiths and Jensen, 2023; Rossi et al., 2001), pointing to rTMS and tACS as good candidates for memory enhancement.

Revised version: Repetitive transcranial magnetic stimulation (rTMS) and transcranial alternating current stimulation (tACS) are two forms of NIBS widely used to enhance memory performance (Grover et al., 2022; Koch et al., 2018; Wang et al., 2014). rTMS, based on the principle of Faraday, induces depolarization of cortical neuronal assemblies and leads to after-effects that have been linked to changes in synaptic plasticity involving mechanisms of long-term potentiation (LTP) (Huang et al., 2017; Jannati et al., 2023). On the other hand, tACS causes rhythmic fluctuations in neuronal membrane potentials, which can bias spike timing, leading to an entrainment of the neural activity (Wischnewski et al., 2023). In particular, the induction of gamma oscillatory activity has been proposed to play an important role in a type of LTP known as spike timing-dependent plasticity, which depends on a precise temporal delay between the firing of a presynaptic and a postsynaptic neuron (Griffiths and Jensen, 2023). Both LTP and gamma oscillations have a strong link with memory processes such as encoding (Bliss and Collingridge, 1993; Griffiths and Jensen, 2023; Rossi et al., 2001), pointing to rTMS and tACS as good candidates for memory enhancement.

However, despite the increasing expectations, their overall effects are relatively modest, even with prolonged stimulation sessions (Grover et al., 2022; Polanía et al., 2018; Wang et al., 2014). In fact, the behavioral and clinical effects are extremely variable (Corp et al., 2020), especially when used to modulate memory functions (Pabst et al., 2022).

Earlier version: Recent studies explored the possibility of combining these two techniques showing that the simultaneous application of tACS in the range of the gamma frequency band (γtACS) with intermittent theta bust (iTBS), an rTMS protocol known to induce LTP (Huang et al., 2005), enhances the after-effects of iTBS alone on the motor and prefrontal cortices(Guerra et al., 2018; Maiella et al., 2022).

Revised version: Recent studies explored the possibility of combining these two techniques showing that the simultaneous application of tACS in the range of the gamma frequency band (γtACS) with intermittent theta burst (iTBS), an rTMS protocol known to induce LTP (Huang et al., 2005), enhances the after-effects of iTBS alone on the motor and prefrontal cortices (Guerra et al., 2018; Maiella et al., 2022).

Earlier version: Long-term memory formation has been associated with the neural activity of the precuneus (PC). The PC represents one of the main hubs of the default mode network (Cavanna and Trimble, 2006; Cunningham et al., 2017; Jitsuishi and Yamaguchi, 2021) and is involved in the formation of associative, episodic, and autobiographic memories (Bonnì et al., 2015; Brodt et al., 2018; Flanagin et al., 2023). Moreover, it is connected with the temporal lobe and the hippocampus (Cunningham et al., 2017; Dadario and Sughrue, 2023; Tanglay et al., 2022). Hence the PC represents an ideal target for NIBS to modulate memory functions with profound clinical implications such as in the case of Alzheimer’s disease (2).

Revised version: Long-term memory formation has been associated with the neural activity of the precuneus (PC). The PC represents one of the main hubs of the default mode network (Cavanna and Trimble, 2006; Cunningham et al., 2017; Jitsuishi and Yamaguchi, 2021) and is involved in the formation of associative, episodic, and autobiographic memories (Bonnì et al., 2015; Brodt et al., 2018; Flanagin et al., 2023). Moreover, it is connected with the temporal lobe and the hippocampus (Cunningham et al., 2017; Dadario and Sughrue, 2023; Tanglay et al., 2022). Hence, the PC represents an ideal target for NIBS to modulate memory functions with profound clinical implications such as in the case of Alzheimer’s disease (Koch et al., 2025, 2024, 2022, 2018).

Earlier version: Here we hypothesized that γtACS would activate key neural mechanisms underlying memory formation favoring the iTBS-mediated plasticity induction by entraining and synchronizing endogenous gamma oscillations, thereby improving long-term memory. Thus, we reasoned that the dual application of iTBS and γtACS over the PC-hippocampus network could be a promising tool for memory enhancement by enhancing plasticity mechanisms.

Revised version: Here, we hypothesized that γtACS would activate key neural mechanisms underlying memory formation favoring the iTBS-mediated plasticity induction by entraining and synchronizing endogenous gamma oscillations, thereby improving long-term memory. Thus, we reasoned that the dual application of iTBS and γtACS over the PC could be a promising tool for memory enhancement by enhancing plasticity mechanisms.

Results

Earlier version: Dual precuneus iTBS+γtACS improves long-term associative memory

Revised version: Dual precuneus iTBS+γtACS improves long-term associative memory performance

In the first experiment, subjects (n=20) were asked to perform two tasks measuring long and short-term memory performances immediately after three balanced, randomized, and double-blind stimulation conditions with a cross-over design (iTBS+γtACS, iTBS+sham-tACS, sham-iTBS+sham-tACS). Stimulation parameters were personalized and identified using fMRI and TMS-EEG (see Fig. 1, panels A and B). Long-term memory was assessed by the face-name associative task (FNAT), requiring the memorization of 12 faces with corresponding names and occupations, while short-term memory was assessed by the short-term memory binding test (STMB) (see Fig. 1, panel C).

Individual target extraction and experimental design.

Earlier version: (A) Methodological workflow used to extract the individualized target for the neuromodulation protocol used in experiments 1, 3 and 4. Participants underwent fMRI scanning to individualize the stimulation sites and permit neuronavigation. Target individualization was derived by computing a PC functional connectivity profile for each participant, thus obtaining a map of positively correlated voxels, respectively representing the DMN (Panel A, left). The individual stimulation targets were defined as the centroid of the strongest PC activation being on the top of a cortical gyrus and representing the shortest perpendicular path connecting the stimulating TMS coil on the scalp and the cortex (Panel A, center). PC coordinates for each participant are represented in red on an MNI brain template, showing the overlap with the DMN (Panel A, right). (B) Biophysical modeling was computed for each subject acquiring T1w and T2w MRI and using the Simnibs toolbox for the T1w segmentation and the 3D-mesh transformation (Panel B, left). The mean Norm E-field was extracted from a target ROI-sphere (10 mm radius) centered on the individual coordinates of the PC (Panel B, center). The simulated induced electric field is shown for a representative subject produced by iTBS (e-field modeling, left) and tACS (e-field modeling, right). EEG source activity reconstruction induced by the TMS pulse over the precuneus in a representative subject (Panel B, right). Experimental design of experiments 1(C), 2 (D), 3 (E), 4 (F). The effect of simultaneous iTBS+γtACS on memory performances was investigated in experiments 1(C) and 2 (D) through two memory tasks: the face-name associative task (FNAT), which required the memorization of 12 faces with corresponding faces and occupations, and the visual short-term memory binding test (STMB), which consisted in a change detection task. In the main experiment 1 (C), subjects were involved in a cross-over design with different experimental sessions of neuromodulation separated by a washout week. Every session corresponded to a different balanced and randomized stimulation condition (i.e., iTBS+γtACS, iTBS+sham-tACS, sham-iTBS+sham-γtACS) immediately followed by the FNAT learning phase and immediate recall, the STMB and the FNAT delayed recall (15-minute delayed) and recognition.

Revised version: (A) Methodological workflow used to extract the individualized target for the neuromodulation protocol used in experiments 1, 3 and 4. Participants underwent fMRI scanning to individualize the stimulation sites and permit neuronavigation. Target individualization was derived by computing a PC functional connectivity profile for each participant, thus obtaining a map of positively correlated voxels, respectively representing the DMN (Panel A, left). The individual stimulation targets were defined as the centroid of the strongest PC activation being on the top of a cortical gyrus and representing the shortest perpendicular path connecting the stimulating TMS coil on the scalp and the cortex (Panel A, center). PC coordinates for each participant are represented in red on an MNI brain template, showing the overlap with the DMN (Panel A, right). (B) Biophysical modeling was computed for each subject acquiring T1w and T2w MRI and using the Simnibs toolbox for the T1w segmentation and the 3D-mesh transformation (Panel B, left). The mean Norm E-field was extracted from a target ROI-sphere (10 mm radius) centered on the individual coordinates of the PC (Panel B, center). The simulated induced electric field is shown for a representative subject produced by iTBS (e-field modeling, left) and tACS (e-field modeling, right). EEG source activity reconstruction induced by the TMS pulse over the precuneus in a representative subject (Panel B, right). Experimental design of experiments 1(C), 2 (D), 3 (E), 4 (F). The effect of simultaneous iTBS+γtACS on memory performances was investigated in experiments 1(C) and 2 (D) through two memory tasks: the face-name associative task (FNAT), which required the memorization of 12 faces with corresponding faces and occupations, and the visual short-term memory binding test (STMB), which consisted in a change detection task. In the main experiment 1 (C), subjects were involved in a cross-over design with different experimental sessions of neuromodulation separated by a washout week. Every session corresponded to a different balanced and randomized stimulation condition (i.e., iTBS+γtACS, iTBS+sham-tACS, sham-iTBS+sham-γtACS) immediately followed by the FNAT learning phase and immediate recall, the STMB and the FNAT delayed recall (15-minute delayed) and recognition. In experiment 2 (D) subjects were involved in a cross-over design with two balanced and randomized stimulation conditions (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. During the first session (day 1), participants received the neuromodulation protocol and then performed the learning phase and immediate recall FNAT, the STMB and the FNAT delayed recall. In the second session (day 2), participants performed FNAT recall with a 24-hour delay from the neuromodulation protocol, while in the third session (day 7), participants performed FNAT recall and recognition with a 1-week delay. In experiment 3 (E) participants were involved in two randomized and balanced experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. TMS-EEG recordings were performed before (T0), immediately after (T1), and 20 minutes after the neuromodulation (T2). In experiment 4 (D), after the first MRI scanning used for neuronavigation, participants were involved in two randomized and balanced experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. The fMRI scanning was performed before (T0) and immediately after (T1) the neuromodulation protocol. Photographs reported represent the author Michele Maiella performing the task and an example of the face item used in the task taken from the FACES database (Ebner et al., 2010).

Removed in the revised version: In experiment 2 (D) subjects were involved in a cross-over design with two balanced and randomized stimulation conditions (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. During the first session (day 1), participants received the neuromodulation protocol and then performed the learning phase and immediate recall FNAT, the STMB and the FNAT delayed recall. In the second session (day 2), participants performed FNAT recall with a 24-hour delay from the neuromodulation protocol, while in the third session (day 7), participants performed FNAT recall and recognition with a 1-week delay. In experiment 3 (E) participants were involved in two randomized and balanced experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. TMS-EEG recordings were performed before (T0), immediately after (T1), and 20 minutes after the neuromodulation (T2). In experiment 4 (D), after the first MRI scanning used for neuronavigation, participants were involved in two randomized and balanced experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. The fMRI scanning was performed before (T0) and immediately after (T1) the neuromodulation protocol.

Removed in the revised version: Photographs reported represent the author Michele Maiella performing the task and an example of the face item used in the task taken from the FACES database (Ebner et al., 2010).

Earlier version: The different stimulation protocols were all well-tolerated. No one reported significant side effects connected with the neuromodulation protocol applications. We found that dual iTBS+γtACS improved long-term associative memory as compared to the other conditions. Dual iTBS+γtACS increased the performances in recalling the association between face, name and occupation (FNAT accuracy) both for the immediate (F2,38=7.18; p=0.002; η²p=0.274) and the delayed (F2,38=5.86; p=0.006; η²p=0.236) recall performances (Fig. 2, panel A). The sensitivity analysis showed a minimal detectable effect size of η²p=0.215 with 20 participants. Post-hoc analysis showed an advantage for dual iTBS+γtACS as compared to iTBS+sham-tACS condition for the immediate (40.0±21.6% vs. 25.0±17.9%, p=0.002 Bonferroni corrected, corresponding to a relative improvement of 63.79%) and for the delayed recall (34.2±20.4% vs. 24.2±19.8%, p<0.05 Bonferroni corrected, corresponding to a relative improvement of 34.16%), as well as compared to the sham-iTBS+sham-tACS condition for the delayed recall (34.2±20.4% vs. 26.3±17.6% p=0.037 Bonferroni corrected, corresponding to a relative improvement of 34.04%) (Fig. 2, panel A).

Revised version: The different stimulation protocols were all well-tolerated. No one reported significant side effects connected with the neuromodulation protocol applications. We found that dual iTBS+γtACS improved long-term associative memory as compared to the other conditions. Dual iTBS+γtACS increased the performance in recalling the association between face, name and occupation (FNAT accuracy), both for the immediate (F2,38=7.18; p=0.002; η²p=0.274) and the delayed (F2,38=5.86; p=0.006; η²p=0.236) recall performances (Fig. 2, panel A). The sensitivity analysis showed a minimal detectable effect size of η²p=0.215 with 20 participants. Post-hoc analysis showed an advantage for dual iTBS+γtACS as compared to iTBS+sham-tACS condition for the immediate (40.0±21.6% vs. 25.0±17.9%, p=0.002 Bonferroni corrected, corresponding to a relative improvement of 63.79%) and for the delayed recall (34.2±20.4% vs. 24.2±19.8%, p<0.05 Bonferroni corrected, corresponding to a relative improvement of 34.16%), as well as compared to the sham-iTBS+sham-tACS condition for the delayed recall (34.2±20.4% vs. 26.3±17.6% p=0.037 Bonferroni corrected, corresponding to a relative improvement of 34.04%) (Fig. 2, panel A).

Memory performance outcome.

In the Box and Whiskers plots, boxes delimit the lower (Q1) and the upper (Q3) quartile, the whiskers extend to the smallest and largest values, dots represent individual values. In the line graphs, colored dots show individual participant data, lines show the connection between the performances of each participant, black dots and lines correspond to the mean performance values. Grey corresponds to iTBS+sham-tACS, blue to sham-iTBS+sham-tACS and red to iTBS+γtACS. * = p<0.05 (A) Face-name associative task (FNAT) accuracy in immediate (right) and delayed (left) trials resulting from experiment 1. N=20. (B) Short-term memory binding task RTs (left) and accuracy (right) resulting from experiment 1. N=20. (C) FNAT’s long-lasting effect resulted from experiment 2. The results are shown over time (day 1, day 2, day 7). N=10.

Earlier version: The in-depth analysis of the FNAT accuracy investigating the specific contribution of face-name and face-occupation recall reveald that dual iTBS+γtACS increased the performances in the association between face and name (FNAT NAME) delayed recall (F2,38 =3.46; p =0.042; η²p =0.154; iTBS+γtACS vs. sham-iTBS+sham-tACS: 42.9±21.5 % vs. 33.8±19 %; p=0.048 Bonferroni corrected) (Fig. S4, supplementary information). iTBS+sham-tACS did not result in any improvement as compared to sham-iTBS+sham-tACS. No significant effects were found over the recall of the association between face and occupation (FNAT OCCUPATION) (p<0.05). Moreover, we found that the effects of dual iTBS+γtACS stimulation were selective for long-term but not short-term memory, since no effects were found on STMB accuracy and RTs (all ps>0.05) (Fig. 2, panel B). Table S1 (supplementary information) shows Experiment 1 statistical details. Biophysical modeling calculations showed that the iTBS induced a mean e-field of 29.55±5.17 V/m, while γtACS led to a much smaller mean e-field of 0.1±0.55 V/m.

Revised version: The in-depth analysis of the FNAT accuracy investigating the specific contribution of face-name and face-occupation recall revealed that dual iTBS+γtACS increased the performance in the association between face and name (FNAT NAME) delayed recall (F2,38 =3.46; p =0.042; η²p =0.154; iTBS+γtACS vs. sham-iTBS+sham-tACS: 42.9±21.5 % vs. 33.8±19 %; p=0.048 Bonferroni corrected) (Fig. S4, supplementary information). iTBS+sham-tACS did not result in any improvement as compared to sham-iTBS+sham-tACS. No significant effects were found over the recall of the association between face and occupation (FNAT OCCUPATION) (p<0.05). Moreover, we found that the effects of dual iTBS+γtACS stimulation were selective for long-term but not short-term memory, since no effects were found on STMB accuracy and RTs (all ps>0.05) (Fig. 2, panel B). Table S1 (supplementary information) shows Experiment 1 statistical details. Biophysical modeling calculations showed that the iTBS induced a mean e-field of 29.55±5.17 V/m, while γtACS led to a much smaller mean e-field of 0.1±0.55 V/m.

Earlier version: Enhanced Memory is still evident after one week

Revised version: Enhanced memory performance is still evident after one week

To confirm the above-described effects of dual iTBS+γtACS on an independent sample (n=10) and to investigate the time course of memory retention we performed a second experiment. In this experiment, subjects were asked to execute the delayed FNAT cued recall 24 hours and a 1-week after the neuromodulatory protocols. As for experiment 1, we confirmed that dual iTBS+γtACS improved immediate recall at TOTAL FNAT as compared to iTBS+sham-tACS (F1,9=7.31; p=0.024; η²p=0.448; 26.7±10.2% vs. 17.5±8.3%; p=0.024 corresponding to a relative improvement of 82.50%) (see Table S2 in supplementary information for statistical details). The sensitivity analysis showed a minimal detectable effect size of η²p=0.388 with 10 participants. More importantly, we found that dual iTBS+γtACS exerted a long-lasting effect on long-term associative memory (F1,9=8.433; p=0.017; η²p=0.484) (Fig. 2, panel C). The sensitivity analysis showed a minimal detectable effect size of η²p=0.125 with 10 participants. We observed that the incremental effect on TOTAL FNAT performance was still evident at 24 hours and one week after the neuromodulation (day 1: 25.8±10.7 % vs. 15±9.5 %, t9 =-2.75, p=0.011 corresponding to a relative improvement of 78.33%; day 2: 25±9.6 % vs. 14.2±9.7 %, t9=-2.89, p=0.009 corresponding to a relative improvement of 88.33%; day 7: 23.3±12.9 % vs. 14.2±12.5 %, t9=-2.40, p=0.020 corresponding to a relative improvement of 80%) (see Table S3 in supplementary information). No significant effects were found for STMB accuracy and RTs (all ps>0.05) as for experiment 1 (see table S2 in supplementary information for statistical details).

Precuneus iTBS+γtACS increases gamma oscillatory activity

Earlier version: We then aimed to investigate the effects of the neuromodulation protocol on cortical oscillations (TRSP) and excitability (TEPs), through TMS-EEG recordings (n=14 subjects, 6 of those also participated in experiment 1). We found that cortical oscillations in the gamma band increased after dual iTBS+γtACS as compared to baseline TMS-EEG recordings (stimulation F1,13=5.073; p=0.042; η²p=0.281) (Fig. 3, panel A). Specifically, we observed an increase of the gamma-TRSP after dual iTBS+γtACS as compared to iTBS+sham-tACS in both time points (ΔT1: 0.045±0.09 vs. -0.013±0.06; t13=1.96; p=0.036; ΔT2: 0.026±0.04 vs. -0.013±0.05; t13=1.92; p=0.039) (Fig. 3, panel B). This effect was specific to PC; indeed, no significant effect was found when testing TMS-EEG over the l-PPC (p>0.05). This result was reflected by increased cortical excitability evident after dual iTBS+γtACS but not after iTBS+sham-tACS condition (Fig. 3, panel C). Specifically, we observed a significant TEP amplitude increase over a cluster of parietal and occipital electrodes comparing T0 vs T1 (mean t-value13=-3.29; all ps<0.05) and over POz (t-value13=-4.74; p<0.05) when comparing T0 vs. T2. No significant differences were found comparing T0 vs. T1 and T2 in the iTBS+sham-tACS condition (ps>0.05).

Revised version: We then aimed to investigate the effects of the neuromodulation protocol on cortical oscillations (TRSP) and excitability (TEPs), through TMS-EEG recordings (n=14 subjects, 6 of whom also participated in experiment 1). We found that cortical oscillations in the gamma band increased after dual iTBS+γtACS as compared to baseline TMS-EEG recordings (stimulation F1,13=5.073; p=0.042; η²p=0.281) (Fig. 3, panel A). Specifically, we observed an increase in the gamma-TRSP after dual iTBS+γtACS as compared to iTBS+sham-tACS in both time points (ΔT1: 0.045±0.09 vs. -0.013±0.06; t13=1.96; p=0.036; ΔT2: 0.026±0.04 vs. -0.013±0.05; t13=1.92; p=0.039) (Fig. 3, panel B). This effect was specific to PC; indeed, no significant effect was found when testing TMS-EEG over the l-PPC (p>0.05). This result was reflected by increased cortical excitability evident after dual iTBS+γtACS but not after iTBS+sham-tACS condition (Fig. 3, panel C). Specifically, we observed a significant TEP amplitude increase over a cluster of parietal and occipital electrodes comparing T0 vs T1 (mean t-value13=-3.29; all ps<0.05) and over POz (t-value13=-4.74; p<0.05) when comparing T0 vs. T2. No significant differences were found comparing T0 vs. T1 and T2 in the iTBS+sham-tACS condition (ps>0.05).

Experiment 3 neurophysiological outcome.

(A) Precuneus (PC) oscillatory activity elicited by iTBS+γtACS (up) and iTBS+sham-tACS (down) when testing PC over the three time points (T0, T1, T2 from left to right). (B) Gamma oscillation changes from baseline after iTBS+γtACS (red) and iTBS+sham-tACS (grey). (C) TMS-evoked potential (TEP) produced over the PC when performing TMS-EEG over PC in the two stimulation conditions: iTBS+γtACS (up-left) and iTBS+sham-tACS (up-right) over the three time points (T0, T1, T2). (Down) Topographies and statistical differences in TEPs amplitude after the different stimulation conditions (iTBS+γtACS, left; iTBS+sham-tACS, right) over the three time points (T0, T1, T2, from left to right). N=14; * = p<0.05; bars depict standard error.

Earlier version: Precuneus iTBS+γtACS increases connectivity with the hippocampus

Revised version: iTBS+γtACS increases precuneus connectivity with the hippocampus

Finally, we investigated the effects of the dual iTBS+γtACS protocols on MRI-based resting state functional connectivity (n=16 subjects, 7 of whom participated in experiment 1). ROI-to-ROI analysis revealed a significant difference specifically for the connectivity between the PC and the hippocampi (F3,13= 5.78, p-FDR= 0.029). Post-hoc analysis showed increased connectivity after the stimulation for the dual iTBS+γtACS (t15 =3.67, p-FDR=0.0023; Fig. 4, panel A), but not for the iTBS+sham-tACS. No significant differences were detected among the conditions before stimulation (all p>0.05). The significant nodes that emerged from the ROI-to-ROI analysis (i.e., PC and hippocampi) were considered for the seed-to-voxel analyses, revealing an increased rs-FC between the PC and the orbitofrontal cortex after dual iTBS+γtACS as compared to baseline (OFC; MNI coordinates: x=20, y=16, z=-4; |T15| > 3.20, k=1051; p=0.0003 (Fig. 4, panel B center). Moreover, stronger connectivity with the PC (MNI coordinates: x=30, y=-70, z=46; |T15| > 3.20, k=1499; p=0.000167) and with the posterior cingulate cortex emerged when considering the right hippocampus as seed (MNI coordinates: x=-14, y=-50, z=12; |T15| > 3.20, k=1295; p=0.0001649) (Fig. 4, panel B right). Finally, similar results occurred considering as seed the left hippocampus, with an improved rs-FC with the PC (MNI coordinates: x=30, y=-68, z=44; |T15| > 3.20, k=1525; p=0.000163) and with the posterior cingulate cortex (MNI coordinates: x=20, y=-50, z=12; |T15| > 3.20, k=509; p=0.022) (Fig. 4, panel B left). We also reconstructed the white matter fibers connecting the PC with the temporal lobe in each subject, forming the bilateral Middle Longitudinal Fasciculus (MdLF), using DTI tractography. We found that the structure of the MdLF measured through fractional anisotropy (FA) was related to the increased connectivity between PC and the hippocampi after dual iTBS+γtACS. A significant positive correlation emerged between the FA of the MdLF and the changes in connectivity between PC and hippocampus in the iTBS+γtACS (r=0.73, p=0.001), but not in the iTBS+sham-tACS (r= 0.14, p=0.634) (Fig. 4, panel C).

rsFC changes after iTBS+γtACS and correlation with DTI.

(A) The standard MNI brain (left) and the bar plot (right) display the positive correlation between the PC and the bilateral HIP after the iTBS+γtACS resulted from the ROI-to-ROI analysis. (B) Seed-to-voxel analysis results from each significant ROI (i.e. left HIP, left; PC, center; right HIP, right) are overlaid to a standard MNI brain. (C) MdLF extracted (left); positive correlation between MdLF integrity and functional connectivity changes between the PC and bilateral HIP after iTBS+γtACS (up right); the absence of correlation in the iTBS+sham-tACS condition and functional connectivity (down right). N=16; * = p<0.05; bars depict standard error.

Discussion

Earlier version: Here we demonstrated that the personalized combination of iTBS and γtACS applied over the PC can enhance long-term associative memory formation. Dual stimulation increased the ability to recall the association between faces, names, and occupations as compared to the iTBS alone or sham conditions. Moreover, the increased memory trace was long-lasting, being evident one week after the stimulation. On the other hand, iTBS alone did not result in any memory improvement, in agreement with previous studies showing contradictory effects of this protocol on neuroenhancement (Hamada et al., 2013; Ziemann and Siebner, 2015). Hence, we show that the combination of neuroplasticity induction promoted by iTBS and gamma oscillations entrainment produced by γtACS, is the key to memory boosting. We believe that the novel results reported here may have several implications for treating memory impairments, gamma dysregulation, and network dysfunction in a large variety of conditions.

Revised version: Here, we demonstrated that the combination of iTBS and γtACS applied over the personalized PC can enhance long-term associative memory formation. Dual stimulation increased the ability to recall the association between faces, names, and occupations as compared to the iTBS alone or sham conditions. Moreover, the increased memory trace was long-lasting, being evident one week after the stimulation. On the other hand, iTBS alone did not result in any memory improvement, in agreement with previous studies showing contradictory effects of this protocol on neuroenhancement (Hamada et al., 2013; Ziemann and Siebner, 2015). Hence, we show that the combination of neuroplasticity induction promoted by iTBS and gamma oscillations entrainment produced by γtACS, may be the key to memory boosting. We believe that the novel results reported here may have several implications for treating memory impairments, gamma dysregulation, and network dysfunction in a large variety of conditions.

Earlier version: Our study identifies a central role for the PC in mediating long-term memory formation. Given the existing evidence on TMS propagation and the computation of the Biophysical model with the Efield, we can reasonably assume that the individually identified PC was a mediator of the observed effects (Ridding and Rothwell, 2007). Moreover, we observed specific cortical changes in the posteromedial parietal areas, as evidenced by the whole-brain analysis conducted on TMS-EEG data and the absence of effect on the lateral posterior parietal cortex used as a control condition. The PC has been associated with memory engram formation and retrieval (Brodt et al., 2018; Flanagin et al., 2023; Hebscher et al., 2019). Moreover, it is a prime candidate in processing relations among entities, including semantic concepts, making its activity fundamental for associative memory formation (Fernandino et al., 2022; Rentz et al., 2011; Summerfield et al., 2020). These notions found support in our current findings, showing that neuromodulation of the PC resulted in robust improvement of long-term associative memory.

Revised version: Our results seem to identify the PC as having a role in mediating long-term memory formation. Given the existing evidence on TMS propagation and the computation of the biophysical model with the E-field, we can assume that the individually identified PC was involved in the observed effects (Ridding and Rothwell, 2007). Moreover, we observed specific cortical changes in the posteromedial parietal areas, as evidenced by the whole-brain analysis conducted on TMS-EEG data and the absence of effect on the lateral posterior parietal cortex used as a control condition. The PC has been associated with memory engram formation and retrieval (Brodt et al., 2018; Flanagin et al., 2023; Hebscher et al., 2019). Moreover, it is a prime candidate in processing relations among entities, including semantic concepts, making its activity fundamental for associative memory formation (Fernandino et al., 2022; Rentz et al., 2011; Summerfield et al., 2020). These notions found support in our current findings, showing that neuromodulation of the PC resulted in an improvement of long-term associative memory.

Earlier version: Visual short-term associative memory, measured by STBM performance, was not modulated by any experimental condition. Even if we cannot exclude the possibility that the stimulation could have influenced short-term verbal associative memory, we expected this result since short-term associative memory is known to rely on a distinct frontoparietal network while FNAT, used to investigate long-term associative memory, has already been associated with the neural activity of the PC and the hippocampus (Parra et al., 2014; Rentz et al., 2011). Moreover, we observed a main effect on the ability to concomitantly recall face, name and occupation, while for single association we observed minimal (face-name) or no effect (face-occupation). Forming face-name associations is widely acknowledged to be particularly difficult due to the lack of contextual properties to formulate an associative link (Werheid and Clare, 2007). In contrast, forming an association between a face and other biographical information, such as occupations or hobbies, is easier (McWeeny et al., 1987). We can hypothesize that retaining multiple associations (face-name-occupation) together is even more effortful and cognitively demanding. For these reasons, the stimulation could have promoted effects in a gradient way related to the difficulty of the task. Apart from improving long-term memory, dual iTBS+γtACS promoted a long-term increase in gamma oscillatory activity, as measured by TMS-EEG. Cortical gamma-band oscillations are thought to be primarily generated by parvalbumin-positive GABA-ergic interneurons interacting with excitatory pyramidal cells (Buzsáki and Wang, 2012; Cobb et al., 1995). The repeated activation of these circuits through gamma oscillations fosters LTP formation which is the neurobiological basis of long-term memory (Bikbaev and Manahan-Vaughan, 2008; Debanne and

Revised version: Visual short-term associative memory, measured by STBM performance, was not modulated by any experimental condition. Even if we cannot exclude the possibility that the stimulation could have influenced short-term verbal associative memory, we expected this result since short-term associative memory is known to rely on a distinct frontoparietal network while FNAT, used to investigate long-term associative memory, has already been associated with the neural activity of the PC and the hippocampus (Parra et al., 2014; Rentz et al., 2011). Moreover, we observed a main effect on the ability to concomitantly recall face, name and occupation, while for single association, we observed minimal (face-name) or no effect (face-occupation). Forming face-name associations is widely acknowledged to be particularly difficult due to the lack of contextual properties to formulate an associative link (Werheid and Clare, 2007). In contrast, forming an association between a face and other biographical information, such as occupations or hobbies, is easier (McWeeny et al., 1987). We can hypothesize that retaining multiple associations (face-name-occupation) together is even more effortful and cognitively demanding. For these reasons, the stimulation could have promoted effects in a gradient way related to the difficulty of the task. Apart from improving long-term memory, dual iTBS+γtACS promoted a long-term increase in gamma oscillatory activity, as measured by TMS-EEG. Cortical gamma-band oscillations are thought to be primarily generated by parvalbumin-positive GABA-ergic interneurons interacting with excitatory pyramidal cells (Buzsáki and Wang, 2012; Cobb et al., 1995). The repeated activation of these circuits through gamma oscillations fosters LTP formation, which is the neurobiological basis of long-term memory (Bikbaev and Manahan-Vaughan, 2008; Debanne and Inglebert, 2023; Headley and Weinberger, 2011). This finding is consistent with the idea that increased gamma oscillatory activity may facilitate cortical plasticity, neural communication, and mechanisms of memory encoding and retrieval, thereby ensuring the successful formation of an episodic memory (Griffiths and Jensen, 2023). We hypothesize that γtACS, by entraining and synchronizing ongoing oscillatory gamma activity, may have facilitated the formation of LTP induced by iTBS.

Removed in the revised version: Inglebert, 2023; Headley and Weinberger, 2011). This finding is consistent with the idea that increased gamma oscillatory activity may facilitate cortical plasticity, neural communication, and mechanisms of memory encoding and retrieval, thereby ensuring the successful formation of an episodic memory (Griffiths and Jensen, 2023). We propose that γtACS, by entraining and synchronizing ongoing oscillatory gamma activity, may have facilitated the formation of LTP induced by iTBS.

Earlier version: Furthermore, we found that dual iTBS+γtACS enhanced the connections between the PC and the hippocampus as measured by fMRI connectivity analysis. Effective neural communication driven by gamma oscillations ensures that, during memory formation, incoming information processed in the neocortex activates the relevant cell assemblies in the hippocampus to ensure associative binding (Griffiths and Jensen, 2023). Gamma-band oscillatory activity drives network connectivity allowing for information processing and binding through synchronization between distant cortical areas (Fries, 2005; Jensen et al., 2007; Miltner et al., 1999; Osipova et al., 2006). Hence we argue that PC gamma oscillations would bind relevant stimuli for perceptual representation, while synchronization between cortical and hippocampal neurons would allow the representation to be encoded into the hippocampus (Nyhus and Curran, 2010).

Revised version: Furthermore, we found that dual iTBS+γtACS enhanced the connections between the PC and the hippocampus as measured by fMRI connectivity analysis. Effective neural communication driven by gamma oscillations ensures that, during memory formation, incoming information processed in the neocortex activates the relevant cell assemblies in the hippocampus to ensure associative binding (Griffiths and Jensen, 2023). Gamma-band oscillatory activity drives network connectivity, allowing for information processing and binding through synchronization between distant cortical areas (Fries, 2005; Jensen et al., 2007; Miltner et al., 1999; Osipova et al., 2006). Hence, we argue that PC gamma oscillations could have bound relevant stimuli for perceptual representation, while synchronization between cortical and hippocampal neurons could have allowed the representation to be encoded into the hippocampus (Nyhus and Curran, 2010).

Earlier version: Finally, we found that higher functional connectivity was supported by higher white matter integrity (FA) of MdLF, the main white matter tract connecting the PC with the temporal lobe. Altogether these findings clarify the anatomo-functional network engagement due to PC stimulation, providing further support for the successful engagement of long-range pathways connecting the PC with the temporal lobe underlying memory formation.

Revised version: Finally, our findings seem to suggest that higher functional connectivity was supported by higher white matter integrity (FA) of MdLF, the main white matter tract connecting the PC with the temporal lobe. Altogether, these results help to trace the possible anatomo-functional network engaged by PC stimulation, providing further support for the successful involvement of long-range pathways connecting the PC with the temporal lobe underlying memory formation.

Earlier version: A strength of the current work is the personalization of stimulation parameters that was achieved by targeting the PC spot more connected with the temporal lobe and by selecting the intensity of stimulation for iTBS using TMS-EEG. Hence we argue that personalization of NIBS protocols is a key factor in reducing the inter-individual variability that limit currently the use of common protocols such as TBS and in increasing the overall beneficial effects on cognitive functions.

Revised version: A strength of the current work is the personalization of stimulation parameters that was achieved by targeting the PC spot more connected with the temporal lobe and by selecting the intensity of stimulation for iTBS using TMS-EEG. Hence, we argue that personalization of NIBS protocols is a key factor in reducing the inter-individual variability that currently limits the use of common protocols such as TBS and in increasing the overall beneficial effects on cognitive functions.

Earlier version: Why should the dual application of γtACS+iTBS lead to such a strong boosting effect as compared to iTBS alone? It is well established that tACS alone does not induce relevant after-effects, especially when applied for a short period of 3 minutes such as in the current case (Guerra et al., 2019, 2018). However, during stimulation, tACS alters membrane potential towards depolarization or hyperpolarization in an oscillatory fashion, manipulating the excitability of neurons that become aligned with the introduced electric field, mostly pyramidal cells in layer V (Fröhlich and McCormick, 2010; Guerra et al., 2018; Reato et al., 2013; Schutter and Hortensius, 2011). Notably, these neurons are characterized by intrinsic resonance and neuroplastic activity. In particular, synchronous neural activation in the γ range is considered crucial for the induction of neuroplasticity and memory formation (Buzsáki and Wang, 2012). Gamma oscillations can synchronize the firing of multiple presynaptic neurons so that they exert a stronger depolarizing effect on the target postsynaptic neuron than if they were to fire in isolation (Sjöström et al., 2001). Synchronizing multiple inputs to a postsynaptic neuron enhances the likelihood of LTP. Hence, gamma oscillations are perhaps ideal because they provide a comparatively short window of excitability that ensures all neurons fire in near-perfect unison (Jensen et al., 2007). This provides a long-sought link between gamma oscillations, LTP, and the formation of new memories. For these reasons, γtACS may have provided the ideal neural substrate for the iTBS to exert its plasticity induction mechanisms. This might have occurred by driving neuronal assemblies in a state devoted to inducing long-term changes in cortical networks involved in memory formation. In humans, iTBS has been inspired by protocols used in animal models to elicit LTP in the hippocampus (Huang et al., 2005).

Revised version: Why should the dual application of γtACS+iTBS lead to such a strong boosting effect as compared to iTBS alone? It is well established that tACS alone does not induce relevant after-effects, especially when applied for a short period of 3 minutes, such as in the current case (Guerra et al., 2019, 2018). However, during stimulation, tACS alters membrane potential towards depolarization or hyperpolarization in an oscillatory fashion, manipulating the excitability of neurons that become aligned with the introduced electric field, mostly pyramidal cells in layer V (Fröhlich and McCormick, 2010; Guerra et al., 2018; Reato et al., 2013; Schutter and Hortensius, 2011). Notably, these neurons are characterized by intrinsic resonance and neuroplastic activity. In particular, synchronous neural activation in the γ range is considered crucial for the induction of neuroplasticity and memory formation (Buzsáki and Wang, 2012). Gamma oscillations can synchronize the firing of multiple presynaptic neurons so that they exert a stronger depolarizing effect on the target postsynaptic neuron than if they were to fire in isolation (Sjöström et al., 2001). Synchronizing multiple inputs to a postsynaptic neuron enhances the likelihood of LTP. Hence, gamma oscillations are perhaps ideal because they provide a comparatively short window of excitability that ensures all neurons fire in near-perfect unison (Jensen et al., 2007). This provides a long-sought link between gamma oscillations, LTP, and the formation of new memories. For these reasons, γtACS may have provided the ideal neural substrate for the iTBS to exert its plasticity-inducing mechanisms. This might have occurred by driving neuronal assemblies in a state devoted to inducing long-term changes in cortical networks involved in memory formation. In humans, iTBS has been inspired by protocols used in animal models to elicit LTP in the hippocampus (Huang et al., 2005).

On top of iTBS, γtACS may have synchronized PC neuronal elements, boosting gamma oscillatory activity during the induction of LTP driven by iTBS. Notably, this interaction is rather specific, since previous studies showed that the application of γtACS alone does not induce any after-effects, while only gamma, but not theta or beta tACS, combined with iTBS induces long- lasting changes in cortical excitability (Guerra et al., 2018; Maiella et al., 2022). Moreover, our biophysical modeling calculations showed that γtACS induced a marginal augmentation of the e- field (∼0.15 mV), which is below the amount of induced current needed to induce an effect in superficial cortical neurons (Opitz et al., 2016). Hence, tACS alone cannot account for the observed increase in memory as well as for enhancing cortical activity.

Earlier version: This study has some limitations. Although we studied TMS and tACS propagation through the E-field modeling and observed an increase in the precuneus gamma oscillatory activity, excitability and connectivity with the hippocampi, we cannot exclude that our results might reflect the consequences of stimulating more superficial parietal regions other than the precuneus nor report direct evidence of microscopic changes in the brain after the stimulation. Invasive neurophysiological recordings in humans for this type of study are not feasible due to ethical constraints. Studies on cadavers or rodents would not fully resolve our question due to significant differences between them (i.e. rodents do not have an anatomical correspondence while cadavers have an alterations in electrical conductivity occurring in postmortem tissue). However, further exploration of this aspect in future studies would help in the understanding of γtACS+iTBS effects.

Revised version: This study has some limitations. Although we studied TMS and tACS propagation through the E- field modeling and observed an increase in the precuneus gamma oscillatory activity, excitability and connectivity with the hippocampi, we cannot exclude that our results might reflect the consequences of stimulating more superficial parietal regions other than the precuneus nor report direct evidence of microscopic changes in the brain after the stimulation. Invasive neurophysiological recordings in humans for this type of study are not feasible due to ethical constraints. Studies on cadavers or rodents would not fully resolve our question due to significant differences between them (i.e., rodents do not have an anatomical correspondence, while cadavers have alterations in electrical conductivity occurring in postmortem tissue). However, further exploration of this aspect in future studies would help in the understanding of γtACS+iTBS effects.

Earlier version: We did not assess the effects of γtACS alone. This decision was based on the findings of Guerra et al. (Guerra et al., 2018), who investigated the same protocol and reported no aftereffects. Given the substantial burden of the experimental design on patients and our primary goal of demonstrating an enhancement of effects compared to the standalone iTBS protocol, we decided to leave out this condition. While examining the effects of γtACS alone could help isolate its specific contribution to this target and memory function, extensive research has shown that achieving a cognitive enhancement aftereffect with tACS alone typically requires around 20–25 minutes of stimulation (Grover et al., 2023). We did not study memory functions, gamma oscillations, and synchronization between the PC and hippocampus in the same session due to technical limitations with the current techniques (i.e., fMRI, TMS-EEG). Consequently, these findings do not allow precise inferences regarding the specific mechanisms by which dual iTBS and γtACS of the precuneus modulate learning and memory. Methodological restraints related to the high-frequency protocols employed in this study did not allow for closed-loop implementation. Moreover, we did not test tACS applied at different frequency bands combined with iTBS since our working hypothesis was based on the combination of gamma activity and plasticity induction. Future studies should further investigate the effects of stimulation on distinct memory processes. In particular, stimulation could be applied before retrieval (Rossi et al., 2001), to better elucidate its specific contribution to the observed enhancements in memory performance. Additionally, it would be worth examining whether repeated stimulation - administered both before encoding and before retrieval - could produce a boosting effect. This is especially relevant in light of findings showing that matching the brain state between retrieval and encoding can significantly enhance memory performance (Javadi et al., 2017). Our findings have profound social and clinical implications.

Revised version: We did not assess the effects of γtACS alone. This decision was based on the findings of Guerra et al. (Guerra et al., 2018), who investigated the same protocol and reported no aftereffects. Given the substantial burden of the experimental design on patients and our primary goal of demonstrating an enhancement of effects compared to the standalone iTBS protocol, we decided to leave out this condition. While examining the effects of γtACS alone could help isolate its specific contribution to this target and memory function, extensive research has shown that achieving a cognitive enhancement aftereffect with tACS alone typically requires around 20–25 minutes of stimulation (Grover et al., 2023). We did not study memory functions, gamma oscillations, and synchronization between the PC and hippocampus in the same session due to technical limitations with the current techniques (i.e., fMRI, TMS-EEG). Consequently, these findings do not allow precise inferences regarding the specific mechanisms by which dual iTBS and γtACS of the precuneus modulate learning and memory. Methodological restraints related to the high-frequency protocols employed in this study did not allow for closed-loop implementation. Moreover, we did not test tACS applied at different frequency bands combined with iTBS since our working hypothesis was based on the combination of gamma activity and plasticity induction. Future studies should further investigate the effects of stimulation on distinct memory processes. In particular, stimulation could be applied before retrieval (Rossi et al., 2001) to better elucidate its specific contribution to the observed enhancements in memory performance. Additionally, it would be worth examining whether repeated stimulation, administered both before encoding and before retrieval, could produce a boosting effect. This is especially relevant in light of findings showing that matching the brain state between retrieval and encoding can significantly enhance memory performance (Javadi et al., 2017).

Added in the revised version: Our findings may have profound social and clinical implications

Earlier version: These results may assume relevance in the context of memory impairment, which is an increasing and challenging aspect of the elderly population. The current findings may be relevant in the case of Alzheimer’s disease, where the progressive memory deficit is accompanied by early involvement of DMN with an accumulation of beta-amyloid plaques, neurofibrillary tangles, atrophy, and functional connectivity primarily targeting the PC (Billette et al., 2022; Buckner et al., 2005; Chen et al., 2017; Raichle et al., 2001). In this perspective, the DMN has been recently identified as a new potential therapeutic target for neuromodulation in AD (2). Moreover, by linking together items and concepts, long-term associative memory is essential for the construction of declarative memory and, broadly for learning and everyday functioning. Hence, dual magnetic and electrical transcranial stimulation could be promising in several social and clinical contexts such as learning deficits, autism spectrum, and attention deficit hyperactivity disorders (Mathalon and Sohal, 2015). Moreover, future studies should investigate the effect of the co-stimulation when administrated prior to recall to clarify the specific

Revised version: These results may assume relevance in the context of memory impairment, which is an increasing and challenging aspect of the elderly population. The current findings may be relevant in the case of Alzheimer’s disease, where the progressive memory deficit is accompanied by early involvement of DMN with an accumulation of beta-amyloid plaques, neurofibrillary tangles, atrophy, and functional connectivity primarily targeting the PC (Billette et al., 2022; Buckner et al., 2005; Chen et al., 2017; Raichle et al., 2001). In this perspective, the DMN has been recently identified as a new potential therapeutic target for neuromodulation in AD (Koch et al., 2025, 2024, 2022, 2018). Moreover, by linking together items and concepts, long-term associative memory is essential for the construction of declarative memory and, broadly for learning and everyday functioning. Hence, dual magnetic and electrical transcranial stimulation could be promising in several social and clinical contexts, such as learning deficits, autism spectrum, and attention deficit hyperactivity disorders (Mathalon and Sohal, 2015).

Materials and methods

Experimental design

We conducted four experiments aimed at understanding the effects of dual iTBS+γtACS on associative memory performance, cortical oscillatory activity and reactivity, and functional connectivity.

Added in the revised version: Experiment 1

Earlier version: Experiment 1. We investigated the effect of dual iTBS+γtACS on memory performances through two memory tasks: the face-name associative task (FNAT) and the visual short-term memory binding test (STMB). Participants first underwent MRI scanning to individualize the stimulation sites and permit neuronavigation (Fig. 1, panel A), then they were involved in a double-blind cross-over design with three experimental sessions of neuromodulation separated by a washout week (Fig. 1, panel C). Every session corresponded to a different balanced and randomized stimulation condition (i.e., iTBS+γtACS, iTBS+sham-tACS, sham-iTBS+sham-γtACS) immediately followed by the two memory tasks. γtACS alone was not administrated since a single 3-minute γtACS does not exert any relevant after-effect (Guerra et al., 2018).

Revised version: We investigated the effect of dual iTBS+γtACS on memory performance through two memory tasks: the face-name associative task (FNAT) and the visual short-term memory binding test (STMB). Participants first underwent MRI scanning to individualize the stimulation sites and permit neuronavigation (Fig. 1, panel A), then they were involved in a double-blind cross-over design with three experimental sessions of neuromodulation separated by a washout week (Fig. 1, panel C). Every session corresponded to a different balanced and randomized stimulation condition (i.e., iTBS+γtACS, iTBS+sham-tACS, sham-iTBS+sham-γtACS) immediately followed by the two memory tasks. γtACS alone was not administered since a single 3-minute γtACS does not exert any relevant after-effect (Guerra et al., 2018).

Added in the revised version: Experiment 2

Earlier version: Experiment 2. We deepen the study of memory effects by confirming experiment 1 results and investigating the course of the memory trace over time after the neuromodulation protocol. We simplified the experimental design by excluding one of the control conditions, i.e. ‘sham-iTBS+sham-γtACS’, as we did not observe any effect related to this condition (see ‘results’ paragraph). In particular, participants were involved in a double-blind cross-over design with two balanced and randomized stimulation conditions (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. Participants were required to attend three experimental sessions for each condition. During the first session (day 1), participants received the neuromodulation protocol and then performed FNAT and STMB. In the second session (day 2), participants performed FNAT recall with a 24-hour delay from the neuromodulation protocol, while in the third session (day 7), participants performed FNAT recall and recognition with a 1-week delay (Fig. 1, panel D).

Revised version: We deepen the study of memory effects by confirming experiment 1 results and investigating the course of the memory trace over time after the neuromodulation protocol. We simplified the experimental design by excluding one of the control conditions, i.e., ‘sham- iTBS+sham-γtACS’, as we did not observe any effect related to this condition (see ‘results’ paragraph). In particular, participants were involved in a double-blind cross-over design with two balanced and randomized stimulation conditions (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. Participants were required to attend three experimental sessions for each condition. During the first session (day 1), participants received the neuromodulation protocol and then performed FNAT and STMB. In the second session (day 2), participants performed FNAT recall with a 24-hour delay from the neuromodulation protocol, while in the third session (day 7), participants performed FNAT recall and recognition with a 1-week delay (Fig. 1, panel D).

Added in the revised version: Experiment 3

Earlier version: Experiment 3. We investigated the neurophysiological effects of the neuromodulation protocol through the combined use of TMS and electroencephalography (TMS-EEG). As for experiment 1, participants underwent an MRI scanning to individualize the stimulation sites and permit neuronavigation. Then, they were involved in two balanced, double-blind, randomized experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. TMS-EEG recordings were performed before (T0), immediately after (T1), and 20 minutes after the neuromodulation (T2) (Fig. 1, panel E).

Revised version: We investigated the neurophysiological effects of the neuromodulation protocol through the combined use of TMS and electroencephalography (TMS-EEG). As for experiment 1, participants underwent an MRI scan to individualize the stimulation sites and permit neuronavigation. Then, they were involved in two balanced, double-blind, randomized experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. TMS-EEG recordings were performed before (T0), immediately after (T1), and 20 minutes after the neuromodulation (T2) (Fig. 1, panel E).

Added in the revised version: Experiment 4

Earlier version: Experiment 4. We investigated the effects of the neuromodulation protocol on functional connectivity through fMRI scanning. After the first MRI scanning used for neuronavigation, participants were involved in two balanced, double-blind, randomized experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. The fMRI scanning was performed before (T0) and immediately after (T1) the neuromodulation protocol (Fig. 1, panel F).

Revised version: We investigated the effects of the neuromodulation protocol on functional connectivity through fMRI scanning. After the first MRI scanning used for neuronavigation, participants were involved in two balanced, double-blind, randomized experimental sessions of neuromodulation (i.e., iTBS+γtACS, iTBS+sham-tACS) separated by a washout week. The fMRI scanning was performed before (T0) and immediately after (T1) the neuromodulation protocol (Fig. 1, panel F).

Participants

We recruited young healthy participants who provided written informed consent approved by the Santa Lucia Foundation IRCCS ethical committee (CE/PROG.923) in accordance with the Declaration of Helsinki. See supplementary information for inclusion and exclusion criteria.

In total, 41 subjects were enrolled. 24 participants were involved in experiment 1, in which we tested the effects of the dual iTBS and γtACS on memory functions. Another group of 10 was enrolled in experiment 2 to replicate the results of experiment 1 in an independent sample and to investigate the long-term effect of the dual stimulation. 16 participants were recruited in experiment 3, of which 10 had already participated in experiment 1, to examine the effects on cortical oscillatory activity and cortical reactivity. Finally, 18 participants were involved in experiment 4, of which 12 had already participated in experiment 1, to investigate functional connectivity through fMRI.

Eight participants (i.e., n=4 from experiment 1, n=2 from experiment 3, n=2 from experiment 4) voluntarily withdrew before completing all the experimental sessions. We analyzed at all time points in n=20 subjects for experiment 1, n=10 for experiment 2, n=14 for experiment 3, and n=16 for experiment 4 (see Table 1 for demographic characteristics). See supplementary information for sample size estimation.

iTBS+γtACS neuromodulation protocol

The co-stimulation consisted of a combination of iTBS and tACS delivered in the gamma band (γtACS) based on previous studies (Guerra et al., 2018; Maiella et al., 2022). The neuromodulation protocol was delivered on the individual PC, based on the individual resting state structural and functional MRI (see ‘MRI data acquisition and preprocessing’ paragraph), and targeted with a stereotaxic neuronavigation system (SofTaxic, E.M.S., Bologna s.r.l.). The active tACS electrode (anode) was placed on the scalp with the iTBS coil above it, and the other tACS electrode (cathode) was placed over the shoulder’s right muscle (Fig. 1, panel C).

Earlier version: tACS was delivered through a Brainstim multifunctional system for low-intensity transcranial electrical stimulation (E.M.S., Bologna s.r.l.) and saline-soaked sponge electrodes (7×5 cm2). γtACS sinusoid frequency wave was set at 70 Hz with an intensity of 1mA for a total duration of 190s. iTBS was delivered through a MagStim Rapid2 magnetic stimulator (Magstim Company, Whitland, Wales, UK) delivering a biphasic waveform pulse (pulse width ∼ 0.1 ms) connected to a figure-of-eight coil (70 mm). iTBS consisted of ten bursts of three pulses at 50 Hz lasting 2 s, repeated every 10 s with an 8 s pause between consecutive trains, for a total of 600 pulses total lasting 190 s (Huang et al., 2005). Biophysical modeling and E-field calculation was conducted to control for the e-field distribution (see supplementary information). Stimulation parameters were personalized and identified using fMRI, TMS-EEG and electromyography (see supplementary information for details).

Revised version: tACS was delivered through a Brainstim multifunctional system for low-intensity transcranial electrical stimulation (E.M.S., Bologna s.r.l.) and saline-soaked sponge electrodes (7×5 cm2). γtACS sinusoid frequency wave was set at 70 Hz with an intensity of 1mA for a total duration of 190s. iTBS was delivered through a MagStim Rapid2 magnetic stimulator (Magstim Company, Whitland, Wales, UK) delivering a biphasic waveform pulse (pulse width ∼ 0.1 ms) connected to a figure-of-eight coil (70 mm). iTBS consisted of ten bursts of three pulses at 50 Hz lasting 2 s, repeated every 10 s with an 8 s pause between consecutive trains, for a total of 600 pulses lasting 190 s (Huang et al., 2005). Biophysical modeling and E-field calculation were conducted to control for the E-field distribution (see supplementary information). Stimulation parameters were personalized and identified using fMRI, TMS-EEG and electromyography (see supplementary information for details).

Earlier version: No ramp-up and no ramp-down were programmed for the stimulation. Sham stimulation conditions were implemented to control the individual contribution of the techniques and the placebo effect. tACS sham conditions were implemented applying only a 2 s ramp up and 2 s ramp down at the beginning and at the end of stimulation, to give the participant the feeling of real stimulation. Sham-iTBS conditions were implemented by adding a wood layer under the coil(Sandrini et al., 2020). See supplementary information for randomization and masking information.

Revised version: No ramp-up and no ramp-down were programmed for the stimulation. Sham stimulation conditions were implemented to control the individual contribution of the techniques and the placebo effect. tACS sham conditions were implemented by applying only a 2s ramp-up and 2s ramp-down at the beginning and the end of stimulation, to give the participant the feeling of real stimulation. Sham-iTBS conditions were implemented by adding a wood layer under the coil(Sandrini et al., 2020). See supplementary information for randomization and masking information.

Memory tasks

Face-name associative task

In experiments 1 and 2, just after the neuromodulation protocol, the participants underwent the FNAT, a cross-modal associative memory test that requires participants to pair pictures of unfamiliar faces with common first names and occupations (Papp et al., 2014). See supplementary information for details about the task construction.

Earlier version: The task consisted of a learning phase followed by a recall and recognition trial. During the learning phase, 12 faces associated with a name and an occupation were shown for 8 seconds on a computer screen. Participants were instructed to read aloud the name and the occupation, to ensure their attention was focused on the items, and to memorize the face-name-occupation association. Right after the learning phase, there was an immediate cued recall phase, in which the participants were asked to recall the name and the occupation of each face previously shown while observing the faces again for 8 seconds each. The delayed cued recall consisted of showing the same faces and asking participants to say each name and occupation. Finally, during the recognition trial, participants were asked to recognize the studied face from a distractor face matched for age and sex. Then, for those names and/or occupations not recollected, they had to identify the correct name and/or occupation among three alternatives each. The distractors were a novel name/occupation and a name/occupation associated with a different face. We considered a correct association when a subject was able to recall all the information for each item (i.e. face, name and occupation), resulting in a total of 36 items to learn and associate. To further investigate the effect on FNAT we also computed a partial recall score accounting for those items where subjects correctly matched only names with faces (FNAT NAME) and only occupations with faces (FNAT OCCUPATION). See supplementary information for score details. In experiment 1, the learning phase was followed by an immediate cued recall and a 15-minute delayed cued recall with recognition (Fig. S2, panel A in supplementary information) while in experiment 2 the learning phase was followed by an immediate cued recall and a 15-minute delayed recall on day 1, a 24-hour delayed cued recall on day 2, and a 1-week delayed cued recall with recognition on day 7 (Fig. S2, panel B in supplementary information). The recognition trial was administered just on the third session to avoid a learning bias over day 2 and day 7 cued recall.

Revised version: The task consisted of a learning phase followed by a recall and recognition trial. During the learning phase, 12 faces associated with a name and an occupation were shown for 8 seconds on a computer screen. Participants were instructed to read aloud the name and the occupation, to ensure their attention was focused on the items, and to memorize the face-name-occupation association. Right after the learning phase, there was an immediate cued recall phase, in which the participants were asked to recall the name and the occupation of each face previously shown while observing the faces again for 8 seconds each. The delayed cued recall consisted of showing the same faces and asking participants to say each name and occupation. Finally, during the recognition trial, participants were asked to recognize the studied face from a distractor face matched for age and sex. Then, for those names and/or occupations not recollected, they had to identify the correct name and/or occupation among three alternatives each. The distractors were a novel name/occupation and a name/occupation associated with a different face. We considered a correct association when a subject was able to recall all the information for each item (i.e., face, name and occupation), resulting in a total of 36 items to learn and associate. To further investigate the effect on FNAT, we also computed a partial recall score accounting for those items where subjects correctly matched only names with faces (FNAT NAME) and only occupations with faces (FNAT OCCUPATION). See supplementary information for score details. In experiment 1, the learning phase was followed by an immediate cued recall and a 15-minute delayed cued recall with recognition (Fig. S2, panel A in supplementary information) while in experiment 2 the learning phase was followed by an immediate cued recall and a 15-minute delayed recall on day 1, a 24-hour delayed cued recall on day 2, and a 1-week delayed cued recall with recognition on day 7 (Fig. S2, panel B in supplementary information). The recognition trial was administered just in the third session to avoid a learning bias over day 2 and day 7 cued recall.

The Short-Term Memory Binding test

Earlier version: This test was run between the FNAT immediate cued recall and the 15-minute delayed cued recall. STMB is a recognition task relying on a change detection paradigm elaborated by Parra and colleagues (Parra et al., 2014). Participants were instructed to remember visual arrays of three black shapes (shape only condition, Fig. S3, panel A in supplementary information) or colored shapes (shape-color binding condition, Fig. S3, panel B in supplementary information) presented for 2 seconds (study phase). After a 1-second delay where a blank screen is presented (retention interval), a display with the same or different items appears in new random locations on the screen (test phase). Participants were asked to press the “1” button on the keyboard if the items shown in the study and test phase were different (50% of the trials) or to press “2” if they were the same. A total of 32 randomized trials were presented for each condition. Conditions were counterbalanced across participants. Before starting the test, each participant underwent a perception trial where the two arrays of shapes were presented on the same screen, to exclude perceptual deficit and to train participants at the keyboard answers (Fig. S3, panel C in supplementary information). The test was administered through E-prime 2.0 (Psychology Software Tools, Pittsburgh, USA) which recorded accuracy and reaction times (RTs). We derived two accuracy indexes, respectively for the shape-only condition and the shape-color binding condition (total score of each condition/32×100).

Revised version: This test was run between the FNAT immediate cued recall and the 15-minute delayed cued recall. STMB is a recognition task relying on a change detection paradigm elaborated by Parra and colleagues (Parra et al., 2014). Participants were instructed to remember visual arrays of three black shapes (shape only condition, Fig. S3, panel A in supplementary information) or colored shapes (shape-color binding condition, Fig. S3, panel B in supplementary information) presented for 2 seconds (study phase). After a 1-second delay where a blank screen is presented (retention interval), a display with the same or different items appears in new random locations on the screen (test phase). Participants were asked to press the “1” button on the keyboard if the items shown in the study and test phase were different (50% of the trials) or to press “2” if they were the same. A total of 32 randomized trials were presented for each condition. Conditions were counterbalanced across participants. Before starting the test, each participant underwent a perception trial where the two arrays of shapes were presented on the same screen, to exclude perceptual deficit and to train participants on keyboard answers (Fig. S3, panel C in supplementary information). The test was administered through E-prime 2.0 (Psychology Software Tools, Pittsburgh, USA), which recorded accuracy and reaction times (RTs). We derived two accuracy indexes, respectively for the shape-only condition and the shape-color binding condition (total score of each condition/32x100).

TMS–EEG data acquisition

Earlier version: Neurophysiological effects in cortical oscillations, excitability, and connectivity were assessed using single-pulse TMS during EEG recordings. During the TMS-EEG assessment, participants sat on a comfortable armchair in a soundproof room, were instructed to fixate a black cross on the wall, and wore in-ear plugs that continuously played a white noise to avoid possible auditory event-related potential responses (Rocchi et al., 2021). TMS-EEG recordings were performed in a resting state to avoid contamination of cortical activity due to underlying memory tasks. The intensity of the white noise was adjusted individually by increasing the volume (always below 90 decibels) until the participant was sure that s/he could no longer hear the TMS-induced click. TMS for EEG recordings was the same stimulator as for the neuromodulation protocol. The stimulated areas were PC, the area of interest, and the left posterior parietal cortex (l-PPC), considered as a control area. The order of the stimulation was counterbalanced across participants. PC and l-PPC were both identified by the individual resting state structural and functional MRI (Fig. 1, panel A; see ‘MRI data acquisition and preprocessing’ paragraph). The coil position was constantly monitored using the Softaxic neuronavigation system (E.M.S. Products, Bologna, Italy) and was differently oriented depending on the area of stimulation so that the direction of current flow in the most effective (second) phase was in the posterior-anterior direction. To target the PC, the coil was positioned with an orientation parallel to the midline, while to target the l-PPC the coil was positioned with an orientation of 15 degrees from the midline (Koch et al., 2018). The single-pulse TMS intensity was set at 100% eSI separately acquired for individual PC and lPPC (see ‘iTBS+γtACS neuromodulation protocol’ paragraph) without the tACS electrode under the coil. Each TMS-EEG block consisted of 90 single pulses with a randomized inter-stimulus interval (ISI) between 2 and 4 s. EEG was continuously recorded from 61 scalp sites positioned according to the 10-20 International System, using TMS-compatible Ag/AgCl pellet electrodes mounted on an elastic cap (BrainAmp; BrainProducts GmbH, Munich, Germany). Additional electrodes were used as a ground and reference: the ground electrode was positioned in Fpz, while the reference was positioned on the tip of the nose. EEG signals were digitized at a sampling rate of 5 kHz. Skin/electrode impedance was maintained below 5 kΩ. See supplementary information for TMS-EEG preprocessing and analysis.

Revised version: Neurophysiological effects in cortical oscillations, excitability, and connectivity were assessed using single-pulse TMS during EEG recordings. During the TMS-EEG assessment, participants sat on a comfortable armchair in a soundproof room, were instructed to fixate on a black cross on the wall, and wore in-ear plugs that continuously played a white noise to avoid possible auditory event-related potential responses (Rocchi et al., 2021). TMS-EEG recordings were performed in a resting state to avoid contamination of cortical activity due to underlying memory tasks. The intensity of the white noise was adjusted individually by increasing the volume (always below 90 decibels) until the participant was sure that s/he could no longer hear the TMS-induced click. TMS for EEG recordings was the same stimulator as for the neuromodulation protocol. The stimulated areas were the PC, the area of interest, and the left posterior parietal cortex (l-PPC), considered as a control area. The order of the stimulation was counterbalanced across participants. PC and l-PPC were both identified by the individual resting state structural and functional MRI (Fig. 1, panel A; see ‘MRI data acquisition and preprocessing’ paragraph). The coil position was constantly monitored using the Softaxic neuronavigation system (E.M.S. Products, Bologna, Italy) and was differently oriented depending on the area of stimulation so that the direction of current flow in the most effective (second) phase was in the posterior-anterior direction. To target the PC, the coil was positioned with an orientation parallel to the midline, while to target the l-PPC, the coil was positioned with an orientation of 15 degrees from the midline (Koch et al., 2018). The single-pulse TMS intensity was set at 100% eSI, separately acquired for individual PC and lPPC (see ‘iTBS+γtACS neuromodulation protocol’ paragraph) without the tACS electrode under the coil. Each TMS-EEG block consisted of 90 single pulses with a randomized inter-stimulus interval (ISI) between 2 and 4 s. EEG was continuously recorded from 61 scalp sites positioned according to the 10-20 International System, using TMS- compatible Ag/AgCl pellet electrodes mounted on an elastic cap (BrainAmp; BrainProducts GmbH, Munich, Germany). Additional electrodes were used as a ground and reference: the ground electrode was positioned in Fpz, while the reference was positioned on the tip of the nose. EEG signals were digitized at a sampling rate of 5 kHz. Skin/electrode impedance was maintained below 5 kΩ. See supplementary information for TMS-EEG preprocessing and analysis.

MRI data acquisition

MRI data were acquired (1) before experiments 1, 3 and 4, to identify and individualize the stimulation target and (2) before and after the neuromodulation protocol, to assess functional connectivity effects (Fig. 1, Panel C).

Imaging was acquired on a Siemens PRISMA scanner with a 64-channel head-coil (Siemens, USA). During the first acquisition (1), both the structural and functional MRI (fMRI) were run, whereas in the following acquisition (2), only the fMRI was acquired.

Earlier version: The structural imaging was acquired using high-resolution T1-weighted (T1w) anatomical images obtained through a 3D-MPRAGE sequence (TR=2500ms, TE=2 ms, TI=1070ms, flip angle (FA)=8°, thickness=1mm, imaging matrix=240 × 240), and a Diffusion Tensor Images (DTI) sequence Diffusion Tensor Images (DTI) sequence (TR=3400 ms, TE=80ms; 121 directions, b-value=1000 s/mm2, thickness=1.79 mm, gap=1.79 mm, flip angle=90), only acquired before the experiment 3 and 4. The fMRI images were acquired using standard echo-planar blood oxygenation level-dependent (BOLD) imaging (TR=800ms, TE=3 ms, flip angle (FA)=52°, thickness=2.4mm, gap=2.4mm). Subjects were instructed to not focus their thoughts on any particular topic, to not cross their arms or legs, and to keep their eyes open.

Revised version: The structural imaging was acquired using high-resolution T1-weighted (T1w) anatomical images obtained through a 3D-MPRAGE sequence (TR=2500ms, TE=2 ms, TI=1070ms, flip angle (FA)=8°, thickness=1mm, imaging matrix=240 × 240), and a Diffusion Tensor Images (DTI) sequence Diffusion Tensor Images (DTI) sequence (TR=3400 ms, TE=80ms; 121 directions, b- value=1000 s/mm2, thickness=1.79 mm, gap=1.79 mm, flip angle=90), only acquired before the experiment 3 and 4. The fMRI images were acquired using standard echo-planar blood oxygenation level-dependent (BOLD) imaging (TR=800ms, TE=3 ms, flip angle (FA)=52°, thickness=2.4mm, gap=2.4mm). Subjects were instructed not to focus their thoughts on any particular topic, not to cross their arms or legs, and to keep their eyes open.

For the individualization of the stimulation sites (Fig. 1, panel A), the functional region of interest (ROI) representing the PC node of the DMN (for experiments 1, 3 and 4) and the l-PPC node of the fronto-parietal network (FPN) (for experiment 3) was derived from the Harvard-Oxford atlas available in CONN (Whitfield-Gabrieli and Nieto-Castanon, 2012). A seed-to-voxel correlation map was computed for each participant, thus obtaining a map of positively and negatively correlated voxels, respectively representing the DMN and FPN. An investigator expert in MRI checked both resting-state functional connectivity (rs-FC) maps and structural MRI data (i.e., T1- weighted images) to identify individual hotspots based on ad-hoc criteria. In particular, the stimulation sites were defined as the ones closest to the local maxima of the rs-FC cluster identified as DMN-PC and FPN-PPC being on the top of a cortical gyrus and representing the shortest perpendicular path connecting the stimulating TMS coil on the scalp and the cortex. Based on best judgment, the resulting set of coordinates was picked as the individual stimulation site. The individual set of coordinates was then transformed using a non-linear transformation to reconstruct the targets in individual brain spaces. Lastly, to ensure a coherent intrasession and intersession stimulation (Cocchi and Zalesky, 2018; Fitzgerald et al., 2009), the individualized targets were marked in the subject’s anatomical MRI and loaded into our neuronavigation system. Both personalized points were used for the TMS-EEG assessment, whereas only the personalized PC was used for the neuromodulation protocol in experiments 1, 3 and 4. See supplementary information for fMRI and DTI data preprocessing.

Statistical analysis

Memory performance and local oscillatory changes data were analyzed through SPSS v22 (IBM, Armonk, NY). Before undergoing parametric or non-parametric statistical procedures, the assumption of normality distribution of data residuals was assessed with Shapiro-Wilk test. The assumption of sphericity was tested with Mauchly test, if this test was significant, we used the Huynh-Feldt correction. The level of significance was set at α=0.05.

To assess memory performances in experiment 1, we run repeated-measures ANOVAs with stimulation condition as a within-subject factor (i.e., iTBS+γtACS; iTBS+sham-tACS; sham- iTBS+sham-tACS) for each dependent variable. In detail, a separate ANOVA was conducted for every FNAT measure (i.e., name, occupation, total) on each memory process (i.e., immediate recall, delayed recall, and recognition). Separate ANOVAs were also conducted for STBM accuracy and RTs both for the shape and the binding condition. Post-hoc comparisons were performed with paired t-tests corrected with Bonferroni method.

Earlier version: To confirm the results obtained from experiment 1 and to investigate the long-lasting effect on long-term memory we performed experiment 2 where we first ran a repeated-measures ANOVA with stimulation condition as a within-subject factor (i.e., iTBS+γtACS, iTBS+sham-tACS) for the immediate recall as for experiment 1. Then a two-way repeated measure ANOVA was conducted for FNAT delayed recall performance where stimulation condition (i.e., iTBS+γtACS, iTBS+sham-tACS) and time (i.e., day1, day2, day7) were considered as within-subject factors. Following the results obtained in this analysis (see ‘results’ paragraph), we conducted exploratory paired t-tests aimed at investigating the effect of stimulation conditions on each time point.

Revised version: To confirm the results obtained from experiment 1 and to investigate the long-lasting effect on long-term memory, we performed experiment 2, where we first ran a repeated-measures ANOVA with stimulation condition as a within-subject factor (i.e., iTBS+γtACS, iTBS+sham-tACS) for the immediate recall, as in experiment 1. Then a two-way repeated measure ANOVA was conducted for FNAT delayed recall performance where stimulation condition (i.e., iTBS+γtACS, iTBS+sham- tACS) and time (i.e., day1, day2, day7) were considered as within-subject factors. Following the results obtained in this analysis (see ‘results’ paragraph), we conducted exploratory paired t-tests aimed at investigating the effect of stimulation conditions on each time point.

Earlier version: Since we do not have an a priori effect size for experiment 1 and 2, we performed a sensitivity power analysis to ensure that these experiments were able to detect the minimum effect size with 80% power and alpha level of 0.05.

Revised version: Since we do not have an a priori effect size for experiments 1 and 2, we performed a sensitivity power analysis to ensure that these experiments were able to detect the minimum effect size with 80% power and alpha level of 0.05.

To assess cortical oscillation changes in experiment 3, we computed repeated-measures ANOVAs with stimulation condition (iTBS+γtACS, iTBS+sham-tACS) and time T1-T0, T2-T0 (ΔT1, ΔT2) as within factors for each frequency band. Also, in this case, paired t-tests were conducted to investigate the effect of stimulation conditions on each time point. A one-tail hypothesis was considered given the precise hypothesis of gamma increase deriving from our previous work (Maiella et al., 2022).

Earlier version: To assess cortical excitability effects, we used multiple paired t-tests. The analysis was run in the four time windows of the TEP waveform (i.e. 10-30 ms, 21-60 ms, 61-120 ms, 121-150 ms after TMS; see paragraph ‘Cortical excitability’ in ‘TMS-EEG data acquisition and preprocessing’) comparing T1 and T2 to T0 in each stimulation condition. To reduce the occurrence of type I errors, we used the Monte Carlo method, which computes the estimates of the significance probabilities from two surrogate distributions constructed by randomly permuting the two original conditions data 3000 times. Moreover, p-values were corrected with the false discovery rate (FDR) method considering the number of electrodes and were considered significant when at least 10 successive t-tests reached the significance threshold (p<0.05). The same statistical analysis was performed to assess changes in source activation. The time series extracted and analyzed from the ROIs (see paragraph ‘source activation’ in ‘TMS-EEG data acquisition and preprocessing’) ranged from -50 to 150 ms from the TMS pulse. As for excitability analysis, we avoided the occurrence of type I errors by implementing Monte Carlo permutation and considering a significant result when at least 10 successive t-tests reached the significance threshold (p<0.05).

Revised version: To assess cortical excitability effects, we used multiple paired t-tests. The analysis was run in the four time windows of the TEP waveform (i.e., 10-30 ms, 21-60 ms, 61-120 ms, 121-150 ms after TMS; see paragraph ‘Cortical excitability’ in ‘TMS-EEG data acquisition and preprocessing’), comparing T1 and T2 to T0 in each stimulation condition. To reduce the occurrence of type I errors, we used the Monte Carlo method, which computes the estimates of the significance probabilities from two surrogate distributions constructed by randomly permuting the two original conditions’ data 3000 times. Moreover, p-values were corrected with the false discovery rate (FDR) method considering the number of electrodes and were considered significant when at least 10 successive t-tests reached the significance threshold (pLJ<LJ0.05). The same statistical analysis was performed to assess changes in source activation. The time series extracted and analyzed from the ROIs (see paragraph ‘source activation’ in ‘TMS-EEG data acquisition and preprocessing’) ranged from -50 to 150 ms from the TMS pulse. As for excitability analysis, we avoided the occurrence of type I errors by implementing Monte Carlo permutation and considering a significant result when at least 10 successive t-tests reached the significance threshold (pLJ<LJ0.05).

Earlier version: To test resting-state functional connectivity (rsFC) modulation we used a general linear model (GLM). The statistical analyses were carried out using the CONN (v.20b) toolbox and Matlab 2018b software (Mathworks, MA, USA) and were performed considering stimulation conditions (i.e., iTBS+γtACS and iTBS+sham tACS) and time points (i.e., pre and post) as factors. In particular, we performed a ROI-to-ROI analysis focusing on 6 specific ROIs, based on our hypothesis: left and right PPC, medial prefrontal cortex (mPFC), PC, left and right hippocampus, chosen based on Harvard-Oxford (cortical and subcortical) atlas (www.fmrib.ox.ac.uk/fsl/)(Desikan et al., 2006). Specifically, rsFC changes were calculated by computing the Pearson correlation coefficient between the average time series extracted from each individual ROI. Multiple comparisons between ROIs were then performed using a two-sided contrast with a cluster level of p<0.05 FDR-corrected. In addition, the significant nodes were considered as seeds for the voxel-wise analysis, specifically performed comparing the pre and post iTBS+γtACS. Temporal correlations were calculated between these seeds and all other voxels in the brain. Results were computed by applying a cluster-level threshold of p<0.05 FDR-corrected.

Revised version: To test resting-state functional connectivity (rsFC) modulation, we used a general linear model (GLM). The statistical analyses were carried out using the CONN (v.20b) toolbox and Matlab 2018b software (Mathworks, MA, USA) and were performed considering stimulation conditions (i.e., iTBS+γtACS and iTBS+sham tACS) and time points (i.e., pre and post) as factors. In particular, we performed a ROI-to-ROI analysis focusing on 6 specific ROIs, based on our hypothesis: left and right PPC, medial prefrontal cortex (mPFC), PC, left and right hippocampus, chosen based on Harvard-Oxford (cortical and subcortical) atlas (www.fmrib.ox.ac.uk/fsl/)(Desikan et al., 2006). Specifically, rsFC changes were calculated by computing the Pearson correlation coefficient between the average time series extracted from each individual ROI. Multiple comparisons between ROIs were then performed using a two-sided contrast with a cluster level of p<0.05 FDR-corrected. In addition, the significant nodes were considered as seeds for the voxel-wise analysis, specifically performed by comparing the pre and post iTBS+γtACS. Temporal correlations were calculated between these seeds and all other voxels in the brain. Results were computed by applying a cluster-level threshold of p<0.05 FDR- corrected.

The sample size for experiments 3 and 4 was estimated based on a previous study using the same protocol (Maiella et al., 2022). Based on the effect size reported in this work (η2=0.291), our power analysis estimated that a sample size of 14 patients would be necessary to obtain the same effect size with 80% power and an alpha level of 0.05.

To assess linear relationships between the white matter integrity and the functional connectivity outcomes, we perform a bivariate correlation between FA values extracted from the MdLF and the functional connectivity modulation induced by the stimulation assessed through the ROI-to- ROI analysis. Correlations were computed with Pearson’s coefficient (two-tailed).

Data will be available upon rational statement request at https://docs.google.com/spreadsheets/d/1lYQHyXFP2EVT1AwCX-rhndQzPk16Ghb_/edit?usp=sharing&ouid=105875083760241337938&rtpof=true&sd=true.